[1363631480] SERVICE ALERT: HHHHH;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 617 processes
[1363631600] SERVICE ALERT: HHHHH;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 683 processes
[1363631650] HOST ALERT: HHHHH;DOWN;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363631690] SERVICE ALERT: HHHHH;Total Processes;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363631720] HOST ALERT: HHHHH;UP;SOFT;2;NRPE v2.12
[1363631720] SERVICE ALERT: HHHHH;Total Processes;OK;SOFT;3;PROCS OK: 387 processes
[1363631742] SERVICE ALERT: HHHHH;Disk usage;WARNING;HARD;3;DISK WARNING - free space: / 25351 MB (74% inode=98%): /u01 22482 MB (46% inode=97%): /data1 73238 MB (30% inode=99%): /data2 107214 MB (44% inode=99%): /archivelog 38093 MB (99% inode=99%): /redolog 26188 MB (91% inode=99%): /backup1 9988 MB (10% inode=99%): /app01 33309 MB (99% inode=99%):
[1363631780] SERVICE ALERT: HHHHH;PING;CRITICAL;HARD;1;PING CRITICAL - Packet loss = 100%
[1363631842] SERVICE ALERT: HHHHH;Disk usage;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363631842] SERVICE ALERT: HHHHH;CPU Statistics;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363631842] SERVICE ALERT: HHHHH;Swap Usage;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363631881] SERVICE ALERT: HHHHH;Current Users;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363631951] SERVICE ALERT: HHHHH;Disk usage;WARNING;HARD;3;DISK WARNING - free space: / 2814 MB (70% inode=93%): /data1 3840 MB (38% inode=99%): /u01 2714 MB (27% inode=90%): /u05 6523 MB (19% inode=99%): /u03 1667 MB (16% inode=98%): /u06 8888 MB (26% inode=99%): /u02 5248 MB (52% inode=99%): /data2 5386 MB (53% inode=99%): /u04 3829 MB (38% inode=99%): /export/home 2716 MB (91% inode=99%):
[1363632290] SERVICE ALERT: HHHHH;Total Processes;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363632320] SERVICE ALERT: HHHHH;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 658 processes
[1363632320] SERVICE ALERT: HHHHH;Memory;WARNING;SOFT;1;WARNING - 2.0% (41376 kB) free!
[1363632380] SERVICE ALERT: HHHHH;PING;CRITICAL;HARD;1;PING CRITICAL - Packet loss = 100%
[1363632440] SERVICE ALERT: HHHHH;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 657 processes
[1363632440] SERVICE ALERT: HHHHH;Memory;OK;SOFT;2;OK - 2.1% (42256 kB) free.
[1363632441] SERVICE ALERT: HHHHH;Disk usage;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363632441] SERVICE ALERT: HHHHH;CPU Statistics;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363632441] SERVICE ALERT: HHHHH;Swap Usage;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363632481] SERVICE ALERT: HHHHH;Current Users;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363632560] SERVICE ALERT: HHHHH;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 665 processes
[1363632660] HOST ALERT: HHHHH;DOWN;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363632680] SERVICE ALERT: HHHHH;Current Load;WARNING;SOFT;1;WARNING - load average: 13.27, 11.07, 6.80
[1363632730] HOST ALERT: HHHHH;UP;SOFT;2;NRPE v2.12
[1363632730] HOST FLAPPING ALERT: HHHHH;STARTED; Host appears to have started flapping (23.7% change > 20.0% threshold)
[1363632800] SERVICE ALERT: HHHHH;Current Load;WARNING;SOFT;2;WARNING - load average: 16.35, 12.69, 7.94
[1363632890] SERVICE ALERT: HHHHH;Total Processes;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363632920] SERVICE ALERT: HHHHH;Current Load;WARNING;HARD;3;WARNING - load average: 18.73, 14.90, 9.37
[1363632940] SERVICE ALERT: HHHHH;Memory;WARNING;HARD;3;WARNING - 5.6% (1847808 kB) free!
[1363632942] SERVICE ALERT: HHHHH;Disk usage;CRITICAL;HARD;3;DISK CRITICAL - free space: / 25351 MB (74% inode=98%): /u01 22489 MB (47% inode=97%): /data1 73238 MB (30% inode=99%): /data2 107214 MB (44% inode=99%): /archivelog 38093 MB (99% inode=99%): /redolog 26188 MB (91% inode=99%): /backup1 4470 MB (4% inode=99%): /app01 33309 MB (99% inode=99%):
[1363632980] SERVICE ALERT: HHHHH;PING;CRITICAL;HARD;1;PING CRITICAL - Packet loss = 100%
[1363633041] SERVICE ALERT: HHHHH;Disk usage;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363633041] SERVICE ALERT: HHHHH;CPU Statistics;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363633042] SERVICE ALERT: HHHHH;Swap Usage;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363633081] SERVICE ALERT: HHHHH;Current Users;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363633160] SERVICE ALERT: HHHHH;Total Processes;OK;HARD;3;PROCS OK: 462 processes
[1363633450] SERVICE ALERT: HHHHH;Current Load;WARNING;SOFT;1;WARNING - load average: 6.03, 6.44, 5.29
[1363633490] SERVICE ALERT: HHHHH;Total Processes;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363633542] SERVICE ALERT: HHHHH;Disk usage;OK;HARD;3;DISK OK - free space: / 25351 MB (74% inode=98%): /u01 22489 MB (47% inode=97%): /data1 73238 MB (30% inode=99%): /data2 107214 MB (44% inode=99%): /archivelog 38093 MB (99% inode=99%): /redolog 26188 MB (91% inode=99%): /backup1 13375 MB (13% inode=99%): /app01 33309 MB (99% inode=99%):
[1363633542] SERVICE FLAPPING ALERT: HHHHH;Disk usage;STARTED; Service appears to have started flapping (23.9% change >= 20.0% threshold)
[1363633572] SERVICE ALERT: HHHHH;Current Load;WARNING;SOFT;2;WARNING - load average: 5.48, 6.12, 5.31
[1363633580] SERVICE ALERT: HHHHH;PING;CRITICAL;HARD;1;PING CRITICAL - Packet loss = 100%
[1363633641] SERVICE ALERT: HHHHH;Disk usage;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363633641] SERVICE ALERT: HHHHH;CPU Statistics;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363633641] SERVICE ALERT: HHHHH;Swap Usage;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363633682] SERVICE ALERT: HHHHH;Current Users;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363633691] SERVICE ALERT: HHHHH;Current Load;WARNING;HARD;3;WARNING - load average: 5.30, 5.84, 5.30
[1363633740] HOST ALERT: HHHHH;DOWN;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363633750] HOST ALERT: HHHHH;UP;SOFT;2;NRPE v2.12
[1363633760] SERVICE ALERT: HHHHH;Total Processes;WARNING;SOFT;1;PROCS WARNING: 530 processes
[1363633761] SERVICE ALERT: HHHHH;Current Load;WARNING;SOFT;1;WARNING - load average: 8.60, 8.55, 6.66
[1363633887] SERVICE ALERT: HHHHH;Total Processes;OK;SOFT;2;PROCS OK: 441 processes
[1363633887] SERVICE ALERT: HHHHH;Current Load;WARNING;SOFT;2;WARNING - load average: 8.50, 8.55, 6.90
[1363633962] SERVICE ALERT: HHHHH;Memory;WARNING;SOFT;1;WARNING - 10.0% (1639424 kB) free!
[1363634001] SERVICE ALERT: HHHHH;Current Load;WARNING;HARD;3;WARNING - load average: 7.23, 8.11, 6.94
[1363634081] SERVICE ALERT: HHHHH;Memory;WARNING;SOFT;2;WARNING - 9.9% (1624032 kB) free!
[1363634090] SERVICE ALERT: HHHHH;Total Processes;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363634141] SERVICE ALERT: HHHHH;Disk usage;WARNING;SOFT;1;DISK WARNING - free space: / 25351 MB (74% inode=98%): /u01 22489 MB (47% inode=97%): /data1 73238 MB (30% inode=99%): /data2 107214 MB (44% inode=99%): /archivelog 38093 MB (99% inode=99%): /redolog 26188 MB (91% inode=99%): /backup1 9432 MB (9% inode=99%): /app01 33309 MB (99% inode=99%):
[1363634180] SERVICE ALERT: HHHHH;PING;CRITICAL;HARD;1;PING CRITICAL - Packet loss = 100%
[1363634203] SERVICE ALERT: HHHHH;Memory;WARNING;HARD;3;WARNING - 9.9% (1621168 kB) free!
[1363634240] SERVICE ALERT: HHHHH;Memory;WARNING;SOFT;1;WARNING - 2.0% (40904 kB) free!
[1363634241] SERVICE ALERT: HHHHH;Disk usage;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363634241] SERVICE ALERT: HHHHH;CPU Statistics;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363634241] SERVICE ALERT: HHHHH;Swap Usage;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363634260] SERVICE ALERT: HHHHH;Disk usage;WARNING;SOFT;2;DISK WARNING - free space: / 25351 MB (74% inode=98%): /u01 22489 MB (47% inode=97%): /data1 73238 MB (30% inode=99%): /data2 107214 MB (44% inode=99%): /archivelog 38093 MB (99% inode=99%): /redolog 26188 MB (91% inode=99%): /backup1 8697 MB (9% inode=99%): /app01 33309 MB (99% inode=99%):
[1363634281] SERVICE ALERT: HHHHH;Current Users;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363634292] SERVICE ALERT: HHHHH;Current Load;OK;HARD;3;OK - load average: 2.12, 3.39, 4.39
[1363634360] SERVICE ALERT: HHHHH;Memory;OK;SOFT;2;OK - 2.1% (42208 kB) free.
[1363634380] SERVICE ALERT: HHHHH;Disk usage;WARNING;HARD;3;DISK WARNING - free space: / 25351 MB (74% inode=98%): /u01 22489 MB (47% inode=97%): /data1 73238 MB (30% inode=99%): /data2 107214 MB (44% inode=99%): /archivelog 38093 MB (99% inode=99%): /redolog 26188 MB (91% inode=99%): /backup1 7911 MB (8% inode=99%): /app01 33309 MB (99% inode=99%):
[1363634491] SERVICE ALERT: HHHHH;Total Processes;WARNING;SOFT;1;PROCS WARNING: 584 processes
[1363634603] SERVICE ALERT: HHHHH;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 666 processes
[1363634692] SERVICE ALERT: HHHHH;Total Processes;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363634722] SERVICE ALERT: HHHHH;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 652 processes
[1363634782] SERVICE ALERT: HHHHH;PING;CRITICAL;HARD;1;PING CRITICAL - Packet loss = 100%
[1363634843] SERVICE ALERT: HHHHH;Disk usage;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363634843] SERVICE ALERT: HHHHH;CPU Statistics;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363634843] SERVICE ALERT: HHHHH;Swap Usage;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363634883] SERVICE ALERT: HHHHH;Current Users;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363634982] SERVICE ALERT: HHHHH;Disk usage;CRITICAL;HARD;3;DISK CRITICAL - free space: / 25351 MB (74% inode=98%): /u01 22489 MB (47% inode=97%): /data1 73238 MB (30% inode=99%): /data2 107214 MB (44% inode=99%): /archivelog 38093 MB (99% inode=99%): /redolog 26188 MB (91% inode=99%): /backup1 3430 MB (3% inode=99%): /app01 33309 MB (99% inode=99%):
[1363635112] SERVICE ALERT: HHHHH;Memory;OK;HARD;3;OK - 25.8% (8509872 kB) free.
[1363635292] SERVICE ALERT: HHHHH;Total Processes;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363635332] SERVICE ALERT: HHHHH;Memory;CRITICAL;HARD;3;CRITICAL - 4.9% (1616080 kB) free!
[1363635332] SERVICE FLAPPING ALERT: HHHHH;Memory;STARTED; Service appears to have started flapping (22.5% change >= 20.0% threshold)
[1363635362] HOST ALERT: HHHHH;DOWN;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363635382] SERVICE ALERT: HHHHH;PING;CRITICAL;HARD;1;PING CRITICAL - Packet loss = 100%
[1363635432] HOST ALERT: HHHHH;UP;SOFT;2;NRPE v2.12
[1363635432] HOST FLAPPING ALERT: HHHHH;STARTED; Host appears to have started flapping (22.9% change > 20.0% threshold)
[1363635443] SERVICE ALERT: HHHHH;Disk usage;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363635443] SERVICE ALERT: HHHHH;CPU Statistics;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363635443] SERVICE ALERT: HHHHH;Swap Usage;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363635483] SERVICE ALERT: HHHHH;Current Users;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363635582] SERVICE ALERT: HHHHH;Disk usage;OK;HARD;3;DISK OK - free space: / 25350 MB (74% inode=98%): /u01 22488 MB (47% inode=97%): /data1 73238 MB (30% inode=99%): /data2 107214 MB (44% inode=99%): /archivelog 38093 MB (99% inode=99%): /redolog 26188 MB (91% inode=99%): /backup1 14952 MB (15% inode=99%): /app01 33309 MB (99% inode=99%):
[1363635763] SERVICE FLAPPING ALERT: HHHHH;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1363635793] SERVICE ALERT: HHHHH;Current Load;WARNING;SOFT;1;WARNING - load average: 2.95, 9.35, 5.38
[1363635803] SERVICE ALERT: HHHHH;Current Load;OK;HARD;3;OK - load average: 4.34, 4.41, 4.81
[1363635892] SERVICE ALERT: HHHHH;Total Processes;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363635912] SERVICE ALERT: HHHHH;Current Load;OK;SOFT;2;OK - load average: 0.50, 6.34, 4.73
[1363635982] SERVICE ALERT: HHHHH;PING;CRITICAL;HARD;1;PING CRITICAL - Packet loss = 100%
[1363636042] SERVICE ALERT: HHHHH;Disk usage;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363636042] SERVICE ALERT: HHHHH;CPU Statistics;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363636042] SERVICE ALERT: HHHHH;Swap Usage;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363636083] SERVICE ALERT: HHHHH;Current Users;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363636162] SERVICE ALERT: HHHHH;Memory;WARNING;SOFT;1;WARNING - 9.6% (3139472 kB) free!
[1363636290] SERVICE ALERT: HHHHH;Memory;WARNING;SOFT;2;WARNING - 9.5% (3133400 kB) free!
[1363636403] SERVICE ALERT: HHHHH;Memory;WARNING;HARD;3;WARNING - 9.6% (3134608 kB) free!
[1363636492] SERVICE ALERT: HHHHH;Total Processes;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363636532] SERVICE ALERT: HHHHH;Memory;WARNING;HARD;3;WARNING - 5.5% (1811000 kB) free!
[1363636582] SERVICE ALERT: HHHHH;PING;CRITICAL;HARD;1;PING CRITICAL - Packet loss = 100%
[1363636642] SERVICE ALERT: HHHHH;Disk usage;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363636643] SERVICE ALERT: HHHHH;CPU Statistics;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363636643] SERVICE ALERT: HHHHH;Swap Usage;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363636683] SERVICE ALERT: HHHHH;Current Users;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363637092] SERVICE ALERT: HHHHH;Total Processes;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363637182] SERVICE ALERT: HHHHH;PING;CRITICAL;HARD;1;PING CRITICAL - Packet loss = 100%
[1363637242] SERVICE ALERT: HHHHH;Disk usage;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363637242] SERVICE ALERT: HHHHH;CPU Statistics;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363637243] SERVICE ALERT: HHHHH;Swap Usage;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363637284] SERVICE ALERT: HHHHH;Current Users;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363637512] SERVICE ALERT: HHHHH;Disk usage;WARNING;SOFT;1;DISK WARNING - free space: / 32038 MB (87% inode=99%): /home 7533 MB (80% inode=99%): /usr 5810 MB (61% inode=93%): /var 8904 MB (94% inode=99%): /u01 31010 MB (63% inode=98%): /archivelog 102127 MB (99% inode=99%): /data1 767377 MB (74% inode=99%): /redologs 45253 MB (88% inode=99%): /backup1 39379 MB (19% inode=99%): /index1 242359 MB (94% inode=99%): /data2 223976 MB (87% inode=99%):
[1363637563] SERVICE ALERT: HHHHH;Memory;OK;HARD;3;OK - 12.3% (3811656 kB) free.
[1363637633] SERVICE ALERT: HHHHH;Disk usage;WARNING;SOFT;2;DISK WARNING - free space: / 32038 MB (87% inode=99%): /home 7533 MB (80% inode=99%): /usr 5810 MB (61% inode=93%): /var 8904 MB (94% inode=99%): /u01 31010 MB (63% inode=98%): /archivelog 102127 MB (99% inode=99%): /data1 767377 MB (74% inode=99%): /redologs 45253 MB (88% inode=99%): /backup1 39115 MB (19% inode=99%): /index1 242359 MB (94% inode=99%): /data2 223976 MB (87% inode=99%):
[1363637692] SERVICE ALERT: HHHHH;Total Processes;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363637752] SERVICE ALERT: HHHHH;Disk usage;WARNING;HARD;3;DISK WARNING - free space: / 32038 MB (87% inode=99%): /home 7533 MB (80% inode=99%): /usr 5810 MB (61% inode=93%): /var 8904 MB (94% inode=99%): /u01 31010 MB (63% inode=98%): /archivelog 102127 MB (99% inode=99%): /data1 767377 MB (74% inode=99%): /redologs 45253 MB (88% inode=99%): /backup1 39115 MB (19% inode=99%): /index1 242359 MB (94% inode=99%): /data2 223976 MB (87% inode=99%):
[1363637782] SERVICE ALERT: HHHHH;PING;CRITICAL;HARD;1;PING CRITICAL - Packet loss = 100%
[1363637842] SERVICE ALERT: HHHHH;Disk usage;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363637842] SERVICE ALERT: HHHHH;CPU Statistics;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363637842] SERVICE ALERT: HHHHH;Swap Usage;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363637883] SERVICE ALERT: HHHHH;Current Users;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363638295] SERVICE ALERT: HHHHH;Total Processes;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363638324] SERVICE ALERT: HHHHH;Total Processes;OK;HARD;3;PROCS OK: 415 processes
[1363638384] SERVICE ALERT: HHHHH;PING;CRITICAL;HARD;1;PING CRITICAL - Packet loss = 100%
[1363638444] SERVICE ALERT: HHHHH;Disk usage;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363638444] SERVICE ALERT: HHHHH;CPU Statistics;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363638444] SERVICE ALERT: HHHHH;Swap Usage;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363638485] SERVICE ALERT: HHHHH;Current Users;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363638604] HOST FLAPPING ALERT: HHHHH;STOPPED; Host appears to have stopped flapping (3.8% change < 5.0% threshold)
[1363638714] SERVICE ALERT: HHHHH;Total Processes;OK;HARD;3;PROCS OK: 482 processes
[1363638895] SERVICE ALERT: HHHHH;Total Processes;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363638924] SERVICE ALERT: HHHHH;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 694 processes
[1363638955] SERVICE ALERT: HHHHH;Disk usage;CRITICAL;HARD;3;DISK CRITICAL - free space: / 31976 MB (87% inode=99%): /home 3724 MB (39% inode=99%): /var 8911 MB (94% inode=99%): /usr 5810 MB (61% inode=93%): /u01 33762 MB (69% inode=97%): /archivelog 101039 MB (98% inode=99%): /data1 327036 MB (19% inode=99%): /redologs 49825 MB (97% inode=99%): /backup1 35211 MB (9% inode=99%): /index1 139337 MB (27% inode=99%): /data2 284843 MB (22% inode=99%):
[1363638984] SERVICE ALERT: HHHHH;PING;CRITICAL;HARD;1;PING CRITICAL - Packet loss = 100%
[1363639034] SERVICE FLAPPING ALERT: HHHHH;Current Load;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1363639044] SERVICE ALERT: HHHHH;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 638 processes
[1363639044] SERVICE ALERT: HHHHH;Disk usage;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363639044] SERVICE ALERT: HHHHH;CPU Statistics;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363639044] SERVICE ALERT: HHHHH;Swap Usage;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363639085] SERVICE ALERT: HHHHH;Current Users;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363639165] SERVICE ALERT: HHHHH;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 678 processes
[1363639314] SERVICE ALERT: HHHHH;Total Processes;WARNING;SOFT;1;PROCS WARNING: 510 processes
[1363639435] SERVICE ALERT: HHHHH;Total Processes;WARNING;SOFT;2;PROCS WARNING: 511 processes
[1363639495] SERVICE ALERT: HHHHH;Total Processes;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363639554] SERVICE ALERT: HHHHH;Total Processes;WARNING;HARD;3;PROCS WARNING: 511 processes
[1363639555] SERVICE ALERT: HHHHH;Disk usage;WARNING;HARD;3;DISK WARNING - free space: / 31976 MB (87% inode=99%): /home 3724 MB (39% inode=99%): /var 8911 MB (94% inode=99%): /usr 5810 MB (61% inode=93%): /u01 33762 MB (69% inode=97%): /archivelog 101039 MB (98% inode=99%): /data1 327036 MB (19% inode=99%): /redologs 49825 MB (97% inode=99%): /backup1 45797 MB (12% inode=99%): /index1 139337 MB (27% inode=99%): /data2 284843 MB (22% inode=99%):
[1363639584] SERVICE ALERT: HHHHH;PING;CRITICAL;HARD;1;PING CRITICAL - Packet loss = 100%
[1363639644] SERVICE ALERT: HHHHH;Disk usage;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363639644] SERVICE ALERT: HHHHH;CPU Statistics;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363639644] SERVICE ALERT: HHHHH;Swap Usage;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363639685] SERVICE ALERT: HHHHH;Current Users;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363640095] SERVICE ALERT: HHHHH;Total Processes;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363640155] SERVICE ALERT: HHHHH;Disk usage;CRITICAL;HARD;3;DISK CRITICAL - free space: / 31976 MB (87% inode=99%): /home 3724 MB (39% inode=99%): /var 8911 MB (94% inode=99%): /usr 5810 MB (61% inode=93%): /u01 33762 MB (69% inode=97%): /archivelog 101039 MB (98% inode=99%): /data1 327036 MB (19% inode=99%): /redologs 49825 MB (97% inode=99%): /backup1 36824 MB (10% inode=99%): /index1 139337 MB (27% inode=99%): /data2 284843 MB (22% inode=99%):
[1363640186] SERVICE ALERT: HHHHH;PING;CRITICAL;HARD;1;PING CRITICAL - Packet loss = 100%
[1363640244] SERVICE ALERT: HHHHH;Disk usage;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363640244] SERVICE ALERT: HHHHH;CPU Statistics;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363640244] SERVICE ALERT: HHHHH;Swap Usage;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363640284] SERVICE ALERT: HHHHH;Current Users;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363640696] SERVICE ALERT: HHHHH;Total Processes;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363640754] SERVICE ALERT: HHHHH;Total Processes;OK;HARD;3;PROCS OK: 481 processes
[1363640755] SERVICE ALERT: HHHHH;Disk usage;WARNING;HARD;3;DISK WARNING - free space: / 31976 MB (87% inode=99%): /home 3724 MB (39% inode=99%): /var 8911 MB (94% inode=99%): /usr 5810 MB (61% inode=93%): /u01 33761 MB (69% inode=97%): /archivelog 101039 MB (98% inode=99%): /data1 327036 MB (19% inode=99%): /redologs 49825 MB (97% inode=99%): /backup1 40919 MB (11% inode=99%): /index1 139337 MB (27% inode=99%): /data2 284843 MB (22% inode=99%):
[1363640755] SERVICE FLAPPING ALERT: HHHHH;Disk usage;STARTED; Service appears to have started flapping (24.2% change >= 20.0% threshold)
[1363640785] SERVICE ALERT: HHHHH;PING;CRITICAL;HARD;1;PING CRITICAL - Packet loss = 100%
[1363640844] SERVICE ALERT: HHHHH;Disk usage;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363640844] SERVICE ALERT: HHHHH;CPU Statistics;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363640844] SERVICE ALERT: HHHHH;Swap Usage;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363640885] SERVICE ALERT: HHHHH;Current Users;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363641195] SERVICE ALERT: HHHHH;Memory;WARNING;SOFT;1;WARNING - 9.5% (3135632 kB) free!
[1363641295] SERVICE ALERT: HHHHH;Total Processes;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363641314] HOST FLAPPING ALERT: HHHHH;STOPPED; Host appears to have stopped flapping (3.8% change < 5.0% threshold)
[1363641314] SERVICE ALERT: HHHHH;Memory;WARNING;SOFT;2;WARNING - 7.7% (2541832 kB) free!
[1363641315] SERVICE ALERT: HHHHH;Current Load;OK;HARD;3;OK - load average: 2.75, 3.09, 4.55
[1363641354] SERVICE ALERT: HHHHH;Total Processes;WARNING;SOFT;1;PROCS WARNING: 502 processes
[1363641355] SERVICE ALERT: HHHHH;Disk usage;CRITICAL;HARD;3;DISK CRITICAL - free space: / 31976 MB (87% inode=99%): /home 3724 MB (39% inode=99%): /var 8911 MB (94% inode=99%): /usr 5810 MB (61% inode=93%): /u01 33762 MB (69% inode=97%): /archivelog 101039 MB (98% inode=99%): /data1 327036 MB (19% inode=99%): /redologs 49825 MB (97% inode=99%): /backup1 31931 MB (8% inode=99%): /index1 139337 MB (27% inode=99%): /data2 284843 MB (22% inode=99%):
[1363641384] SERVICE ALERT: HHHHH;PING;CRITICAL;HARD;1;PING CRITICAL - Packet loss = 100%
[1363641434] SERVICE ALERT: HHHHH;Memory;OK;SOFT;3;OK - 10.9% (3580992 kB) free.
[1363641444] SERVICE ALERT: HHHHH;Disk usage;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363641444] SERVICE ALERT: HHHHH;CPU Statistics;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363641444] SERVICE ALERT: HHHHH;Swap Usage;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363641475] SERVICE ALERT: HHHHH;Total Processes;WARNING;SOFT;2;PROCS WARNING: 501 processes
[1363641485] SERVICE ALERT: HHHHH;Current Users;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363641595] SERVICE ALERT: HHHHH;Total Processes;WARNING;HARD;3;PROCS WARNING: 599 processes
[1363641599] SERVICE ALERT: HHHHH;Total Processes;WARNING;HARD;3;PROCS WARNING: 505 processes
[1363641599] SERVICE FLAPPING ALERT: HHHHH;Total Processes;STARTED; Service appears to have started flapping (23.9% change >= 20.0% threshold)
[1363641895] SERVICE ALERT: HHHHH;Total Processes;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363641984] SERVICE ALERT: HHHHH;PING;CRITICAL;HARD;1;PING CRITICAL - Packet loss = 100%
[1363642044] SERVICE ALERT: HHHHH;Disk usage;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363642044] SERVICE ALERT: HHHHH;CPU Statistics;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363642044] SERVICE ALERT: HHHHH;Swap Usage;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363642085] SERVICE ALERT: HHHHH;Current Users;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363642195] SERVICE ALERT: HHHHH;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 602 processes
[1363642198] SERVICE ALERT: HHHHH;Total Processes;OK;HARD;3;PROCS OK: 499 processes
[1363642496] SERVICE ALERT: HHHHH;Total Processes;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363642584] SERVICE ALERT: HHHHH;PING;CRITICAL;HARD;1;PING CRITICAL - Packet loss = 100%
[1363642644] SERVICE ALERT: HHHHH;Disk usage;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363642644] SERVICE ALERT: HHHHH;CPU Statistics;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363642644] SERVICE ALERT: HHHHH;Swap Usage;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363642684] SERVICE ALERT: HHHHH;Current Users;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363642795] SERVICE ALERT: HHHHH;Total Processes;WARNING;HARD;3;PROCS WARNING: 600 processes
[1363642944] SERVICE ALERT: HHHHH;CPU Statistics;WARNING;SOFT;1;CPU STATISTICS WARNING : user=17% system=2% iowait=41% idle=40%
[1363643065] SERVICE ALERT: HHHHH;CPU Statistics;OK;SOFT;2;CPU STATISTICS OK : user=16% system=5% iowait=22% idle=56%
[1363643094] SERVICE ALERT: HHHHH;Total Processes;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363643184] SERVICE ALERT: HHHHH;PING;CRITICAL;HARD;1;PING CRITICAL - Packet loss = 100%
[1363643244] SERVICE ALERT: HHHHH;Disk usage;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363643244] SERVICE ALERT: HHHHH;CPU Statistics;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363643244] SERVICE ALERT: HHHHH;Swap Usage;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363643284] SERVICE ALERT: HHHHH;Current Users;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363643694] SERVICE ALERT: HHHHH;Total Processes;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363643784] SERVICE ALERT: HHHHH;PING;CRITICAL;HARD;1;PING CRITICAL - Packet loss = 100%
[1363643844] SERVICE ALERT: HHHHH;Disk usage;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363643844] SERVICE ALERT: HHHHH;CPU Statistics;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363643844] SERVICE ALERT: HHHHH;Swap Usage;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363643885] SERVICE ALERT: HHHHH;Current Users;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363643964] SERVICE ALERT: HHHHH;Total Processes;OK;HARD;3;PROCS OK: 447 processes
[1363644294] SERVICE ALERT: HHHHH;Total Processes;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363644295] SERVICE ALERT: HHHHH;Memory;WARNING;SOFT;1;WARNING - 9.7% (3193512 kB) free!
[1363644384] SERVICE ALERT: HHHHH;PING;CRITICAL;HARD;1;PING CRITICAL - Packet loss = 100%
[1363644416] SERVICE ALERT: HHHHH;Memory;OK;SOFT;2;OK - 10.4% (3408272 kB) free.
[1363644444] SERVICE ALERT: HHHHH;CPU Statistics;WARNING;SOFT;1;CPU STATISTICS WARNING : user=3% system=4% iowait=63% idle=30%
[1363644444] SERVICE ALERT: HHHHH;Disk usage;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363644444] SERVICE ALERT: HHHHH;CPU Statistics;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363644444] SERVICE ALERT: HHHHH;Swap Usage;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363644484] SERVICE ALERT: HHHHH;Current Users;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363644564] SERVICE ALERT: HHHHH;CPU Statistics;WARNING;SOFT;2;CPU STATISTICS WARNING : user=1% system=3% iowait=42% idle=54%
[1363644565] SERVICE ALERT: HHHHH;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 671 processes
[1363644595] SERVICE ALERT: HHHHH;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 604 processes
[1363644595] SERVICE FLAPPING ALERT: HHHHH;Total Processes;STARTED; Service appears to have started flapping (23.4% change >= 20.0% threshold)
[1363644685] SERVICE ALERT: HHHHH;CPU Statistics;WARNING;HARD;3;CPU STATISTICS WARNING : user=2% system=4% iowait=34% idle=59%
[1363644689] SERVICE ALERT: HHHHH;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 642 processes
[1363644805] SERVICE ALERT: HHHHH;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 729 processes
[1363644895] SERVICE ALERT: HHHHH;Total Processes;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363644915] SERVICE ALERT: HHHHH;Current Load;WARNING;SOFT;1;WARNING - load average: 6.35, 6.28, 5.20
[1363644984] SERVICE ALERT: HHHHH;PING;CRITICAL;HARD;1;PING CRITICAL - Packet loss = 100%
[1363645034] SERVICE ALERT: HHHHH;Current Load;WARNING;SOFT;2;WARNING - load average: 6.49, 6.32, 5.34
[1363645044] SERVICE ALERT: HHHHH;Disk usage;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363645044] SERVICE ALERT: HHHHH;CPU Statistics;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363645044] SERVICE ALERT: HHHHH;Swap Usage;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363645084] SERVICE ALERT: HHHHH;Current Users;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363645159] SERVICE ALERT: HHHHH;Current Load;WARNING;HARD;3;WARNING - load average: 6.25, 6.33, 5.47
[1363645195] SERVICE ALERT: HHHHH;Total Processes;WARNING;HARD;3;PROCS WARNING: 599 processes
[1363645304] SERVICE ALERT: HHHHH;Current Load;WARNING;SOFT;1;WARNING - load average: 7.40, 6.42, 5.28
[1363645424] SERVICE ALERT: HHHHH;Current Load;WARNING;SOFT;2;WARNING - load average: 7.21, 6.66, 5.51
[1363645494] SERVICE ALERT: HHHHH;Total Processes;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363645544] SERVICE ALERT: HHHHH;Current Load;WARNING;HARD;3;WARNING - load average: 7.13, 6.81, 5.71
[1363645584] SERVICE ALERT: HHHHH;PING;CRITICAL;HARD;1;PING CRITICAL - Packet loss = 100%
[1363645644] SERVICE ALERT: HHHHH;Disk usage;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363645644] SERVICE ALERT: HHHHH;CPU Statistics;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363645644] SERVICE ALERT: HHHHH;Swap Usage;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363645685] SERVICE ALERT: HHHHH;Current Users;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363645885] SERVICE ALERT: HHHHH;CPU Statistics;OK;HARD;3;CPU STATISTICS OK : user=2% system=2% iowait=1% idle=96%
[1363646094] SERVICE ALERT: HHHHH;Total Processes;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363646184] SERVICE ALERT: HHHHH;PING;CRITICAL;HARD;1;PING CRITICAL - Packet loss = 100%
[1363646244] SERVICE ALERT: HHHHH;Disk usage;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363646244] SERVICE ALERT: HHHHH;CPU Statistics;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363646244] SERVICE ALERT: HHHHH;Swap Usage;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363646284] SERVICE ALERT: HHHHH;Current Users;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363646365] SERVICE ALERT: HHHHH;Memory;WARNING;SOFT;1;WARNING - 2.0% (41608 kB) free!
[1363646489] SERVICE ALERT: HHHHH;Memory;WARNING;SOFT;2;WARNING - 2.0% (41584 kB) free!
[1363646605] SERVICE ALERT: HHHHH;Total Processes;OK;HARD;3;PROCS OK: 382 processes
[1363646605] SERVICE ALERT: HHHHH;Memory;WARNING;HARD;3;WARNING - 2.0% (41552 kB) free!
[1363646694] SERVICE ALERT: HHHHH;Total Processes;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363646784] SERVICE ALERT: HHHHH;PING;CRITICAL;HARD;1;PING CRITICAL - Packet loss = 100%
[1363646844] SERVICE ALERT: HHHHH;Disk usage;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363646844] SERVICE ALERT: HHHHH;CPU Statistics;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363646844] SERVICE ALERT: HHHHH;Swap Usage;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363646884] SERVICE ALERT: HHHHH;Current Users;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363646984] SERVICE FLAPPING ALERT: HHHHH;Disk usage;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1363647206] SERVICE ALERT: HHHHH;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 666 processes
[1363647294] SERVICE ALERT: HHHHH;Total Processes;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363647325] SERVICE ALERT: HHHHH;Total Processes;OK;SOFT;2;PROCS OK: 394 processes
[1363647334] SERVICE FLAPPING ALERT: HHHHH;Memory;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1363647344] SERVICE ALERT: HHHHH;Current Load;OK;HARD;3;OK - load average: 4.34, 4.41, 4.78
[1363647384] SERVICE ALERT: HHHHH;PING;CRITICAL;HARD;1;PING CRITICAL - Packet loss = 100%
[1363647444] SERVICE ALERT: HHHHH;Disk usage;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363647444] SERVICE ALERT: HHHHH;CPU Statistics;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363647444] SERVICE ALERT: HHHHH;Swap Usage;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363647484] SERVICE ALERT: HHHHH;Current Users;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363647578] SERVICE ALERT: HHHHH;PING;UNKNOWN;SOFT;1;/usr/sbin/ping -n -s 10.10.1.12 56 5
[1363647694] SERVICE ALERT: HHHHH;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.10 ms
[1363647894] SERVICE ALERT: HHHHH;Total Processes;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363647924] SERVICE ALERT: HHHHH;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 689 processes
[1363647984] SERVICE ALERT: HHHHH;PING;CRITICAL;HARD;1;PING CRITICAL - Packet loss = 100%
[1363648034] SERVICE ALERT: HHHHH;Memory;CRITICAL;SOFT;1;CRITICAL - 4.8% (1584232 kB) free!
[1363648044] SERVICE ALERT: HHHHH;Disk usage;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363648044] SERVICE ALERT: HHHHH;CPU Statistics;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363648044] SERVICE ALERT: HHHHH;Swap Usage;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363648044] SERVICE ALERT: HHHHH;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 677 processes
[1363648086] SERVICE ALERT: HHHHH;Current Users;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363648158] SERVICE ALERT: HHHHH;Memory;WARNING;SOFT;2;WARNING - 5.4% (1764768 kB) free!
[1363648158] SERVICE ALERT: HHHHH;Current Load;OK;HARD;3;OK - load average: 2.40, 2.03, 3.54
[1363648165] SERVICE ALERT: HHHHH;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 624 processes
[1363648195] SERVICE ALERT: HHHHH;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 602 processes
[1363648274] SERVICE ALERT: HHHHH;Memory;WARNING;HARD;3;WARNING - 5.1% (1678488 kB) free!
[1363648494] SERVICE ALERT: HHHHH;Total Processes;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363648584] SERVICE ALERT: HHHHH;PING;CRITICAL;HARD;1;PING CRITICAL - Packet loss = 100%
[1363648644] SERVICE ALERT: HHHHH;Disk usage;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363648644] SERVICE ALERT: HHHHH;CPU Statistics;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363648644] SERVICE ALERT: HHHHH;Swap Usage;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363648684] SERVICE ALERT: HHHHH;Current Users;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363648795] SERVICE ALERT: HHHHH;Total Processes;WARNING;HARD;3;PROCS WARNING: 600 processes
[1363649094] SERVICE ALERT: HHHHH;Total Processes;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363649184] SERVICE ALERT: HHHHH;PING;CRITICAL;HARD;1;PING CRITICAL - Packet loss = 100%
[1363649244] SERVICE ALERT: HHHHH;Disk usage;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363649244] SERVICE ALERT: HHHHH;CPU Statistics;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363649244] SERVICE ALERT: HHHHH;Swap Usage;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363649285] SERVICE ALERT: HHHHH;Current Users;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363649364] SERVICE ALERT: HHHHH;Total Processes;OK;HARD;3;PROCS OK: 414 processes
[1363649694] SERVICE ALERT: HHHHH;Total Processes;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363649734] SERVICE ALERT: HHHHH;Memory;CRITICAL;HARD;3;CRITICAL - 5.0% (1655032 kB) free!
[1363649784] SERVICE ALERT: HHHHH;PING;CRITICAL;HARD;1;PING CRITICAL - Packet loss = 100%
[1363649844] SERVICE ALERT: HHHHH;Disk usage;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363649844] SERVICE ALERT: HHHHH;CPU Statistics;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363649844] SERVICE ALERT: HHHHH;Swap Usage;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363649884] SERVICE ALERT: HHHHH;Current Users;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363649965] SERVICE ALERT: HHHHH;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 664 processes
[1363650091] SERVICE ALERT: HHHHH;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 697 processes
[1363650206] SERVICE ALERT: HHHHH;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 654 processes
[1363650294] SERVICE ALERT: HHHHH;Total Processes;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363650385] SERVICE ALERT: HHHHH;PING;CRITICAL;HARD;1;PING CRITICAL - Packet loss = 100%
[1363650444] SERVICE ALERT: HHHHH;Disk usage;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363650444] SERVICE ALERT: HHHHH;CPU Statistics;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363650444] SERVICE ALERT: HHHHH;Swap Usage;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363650484] SERVICE ALERT: HHHHH;Current Users;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363650674] SERVICE ALERT: HHHHH;Memory;CRITICAL;HARD;3;CRITICAL - 4.8% (1590344 kB) free!
[1363650894] SERVICE ALERT: HHHHH;Total Processes;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363650985] SERVICE ALERT: HHHHH;PING;CRITICAL;HARD;1;PING CRITICAL - Packet loss = 100%
[1363651004] SERVICE ALERT: HHHHH;Memory;CRITICAL;HARD;3;CRITICAL - 4.6% (748792 kB) free!
[1363651044] SERVICE ALERT: HHHHH;Disk usage;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363651044] SERVICE ALERT: HHHHH;CPU Statistics;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363651044] SERVICE ALERT: HHHHH;Swap Usage;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363651084] SERVICE ALERT: HHHHH;Current Users;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363651198] SERVICE ALERT: HHHHH;CPU Statistics;WARNING;SOFT;1;CPU STATISTICS WARNING : user=14% system=7% iowait=36% idle=44%
[1363651274] SERVICE ALERT: HHHHH;Memory;WARNING;HARD;3;WARNING - 5.1% (1685184 kB) free!
[1363651314] SERVICE ALERT: HHHHH;CPU Statistics;OK;SOFT;2;CPU STATISTICS OK : user=4% system=3% iowait=24% idle=69%
[1363651494] SERVICE ALERT: HHHHH;Total Processes;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363651585] SERVICE ALERT: HHHHH;PING;CRITICAL;HARD;1;PING CRITICAL - Packet loss = 100%
[1363651644] SERVICE ALERT: HHHHH;Disk usage;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363651644] SERVICE ALERT: HHHHH;CPU Statistics;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363651644] SERVICE ALERT: HHHHH;Swap Usage;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363651684] SERVICE ALERT: HHHHH;Current Users;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363651795] SERVICE ALERT: HHHHH;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 603 processes
[1363651914] SERVICE ALERT: HHHHH;CPU Statistics;WARNING;SOFT;1;CPU STATISTICS WARNING : user=4% system=2% iowait=70% idle=24%
[1363652005] SERVICE ALERT: HHHHH;Total Processes;OK;HARD;3;PROCS OK: 442 processes
[1363652034] SERVICE ALERT: HHHHH;CPU Statistics;WARNING;SOFT;2;CPU STATISTICS WARNING : user=9% system=2% iowait=65% idle=24%
[1363652094] SERVICE ALERT: HHHHH;Total Processes;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363652154] SERVICE ALERT: HHHHH;CPU Statistics;WARNING;HARD;3;CPU STATISTICS WARNING : user=4% system=3% iowait=59% idle=34%
[1363652185] SERVICE ALERT: HHHHH;PING;CRITICAL;HARD;1;PING CRITICAL - Packet loss = 100%
[1363652244] SERVICE ALERT: HHHHH;Disk usage;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363652244] SERVICE ALERT: HHHHH;CPU Statistics;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363652244] SERVICE ALERT: HHHHH;Swap Usage;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363652285] SERVICE ALERT: HHHHH;Current Users;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363652395] SERVICE ALERT: HHHHH;Total Processes;WARNING;HARD;3;PROCS WARNING: 599 processes
[1363652694] SERVICE ALERT: HHHHH;Total Processes;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363652754] SERVICE FLAPPING ALERT: HHHHH;Disk usage;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1363652785] SERVICE ALERT: HHHHH;PING;CRITICAL;HARD;1;PING CRITICAL - Packet loss = 100%
[1363652844] SERVICE ALERT: HHHHH;Disk usage;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363652844] SERVICE ALERT: HHHHH;CPU Statistics;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363652844] SERVICE ALERT: HHHHH;Swap Usage;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363652885] SERVICE ALERT: HHHHH;Current Users;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363653206] SERVICE ALERT: HHHHH;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 690 processes
[1363653294] SERVICE ALERT: HHHHH;Total Processes;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363653324] SERVICE ALERT: HHHHH;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 647 processes
[1363653354] SERVICE ALERT: HHHHH;Disk usage;OK;HARD;3;DISK OK - free space: / 32038 MB (87% inode=99%): /home 7533 MB (80% inode=99%): /usr 5810 MB (61% inode=93%): /var 8902 MB (94% inode=99%): /u01 31009 MB (63% inode=98%): /archivelog 102127 MB (99% inode=99%): /data1 767377 MB (74% inode=99%): /redologs 45253 MB (88% inode=99%): /backup1 63210 MB (30% inode=99%): /index1 242359 MB (94% inode=99%): /data2 223976 MB (87% inode=99%):
[1363653384] SERVICE ALERT: HHHHH;PING;CRITICAL;HARD;1;PING CRITICAL - Packet loss = 100%
[1363653444] SERVICE ALERT: HHHHH;Disk usage;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363653444] SERVICE ALERT: HHHHH;CPU Statistics;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363653444] SERVICE ALERT: HHHHH;Swap Usage;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363653444] SERVICE ALERT: HHHHH;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 661 processes
[1363653485] SERVICE ALERT: HHHHH;Current Users;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363653596] SERVICE FLAPPING ALERT: HHHHH;Total Processes;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1363653894] SERVICE ALERT: HHHHH;Total Processes;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363653954] SERVICE ALERT: HHHHH;CPU Statistics;OK;HARD;3;CPU STATISTICS OK : user=1% system=1% iowait=14% idle=85%
[1363653985] SERVICE ALERT: HHHHH;PING;CRITICAL;HARD;1;PING CRITICAL - Packet loss = 100%
[1363654044] SERVICE ALERT: HHHHH;Disk usage;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363654044] SERVICE ALERT: HHHHH;CPU Statistics;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363654044] SERVICE ALERT: HHHHH;Swap Usage;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363654085] SERVICE ALERT: HHHHH;Current Users;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363654094] SERVICE ALERT: HHHHH;Current Load;WARNING;SOFT;1;WARNING - load average: 5.83, 6.12, 5.53
[1363654214] SERVICE ALERT: HHHHH;Current Load;WARNING;SOFT;2;WARNING - load average: 5.16, 5.60, 5.40
[1363654334] SERVICE ALERT: HHHHH;Current Load;WARNING;HARD;3;WARNING - load average: 4.32, 5.13, 5.24
[1363654494] SERVICE ALERT: HHHHH;Total Processes;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363654585] SERVICE ALERT: HHHHH;PING;CRITICAL;HARD;1;PING CRITICAL - Packet loss = 100%
[1363654644] SERVICE ALERT: HHHHH;Disk usage;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363654644] SERVICE ALERT: HHHHH;CPU Statistics;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363654644] SERVICE ALERT: HHHHH;Swap Usage;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363654684] SERVICE ALERT: HHHHH;Current Users;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363654795] SERVICE ALERT: HHHHH;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 601 processes
[1363654934] SERVICE ALERT: HHHHH;Current Load;OK;HARD;3;OK - load average: 4.82, 4.74, 4.89
[1363655006] SERVICE ALERT: HHHHH;Memory;OK;HARD;3;OK - 2.6% (52760 kB) free.
[1363655094] SERVICE ALERT: HHHHH;Total Processes;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363655185] SERVICE ALERT: HHHHH;PING;CRITICAL;HARD;1;PING CRITICAL - Packet loss = 100%
[1363655244] SERVICE ALERT: HHHHH;Disk usage;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363655244] SERVICE ALERT: HHHHH;CPU Statistics;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363655244] SERVICE ALERT: HHHHH;Swap Usage;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363655285] SERVICE ALERT: HHHHH;Current Users;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363655474] SERVICE ALERT: HHHHH;Memory;CRITICAL;HARD;3;CRITICAL - 4.6% (1512960 kB) free!
[1363655474] SERVICE FLAPPING ALERT: HHHHH;Memory;STARTED; Service appears to have started flapping (21.4% change >= 20.0% threshold)
[1363655694] SERVICE ALERT: HHHHH;Total Processes;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363655785] SERVICE ALERT: HHHHH;PING;CRITICAL;HARD;1;PING CRITICAL - Packet loss = 100%
[1363655844] SERVICE ALERT: HHHHH;Disk usage;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363655844] SERVICE ALERT: HHHHH;CPU Statistics;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363655844] SERVICE ALERT: HHHHH;Swap Usage;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363655885] SERVICE ALERT: HHHHH;Current Users;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363655995] SERVICE ALERT: HHHHH;Total Processes;WARNING;HARD;3;PROCS WARNING: 600 processes
[1363656294] SERVICE ALERT: HHHHH;Total Processes;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363656385] SERVICE ALERT: HHHHH;PING;CRITICAL;HARD;1;PING CRITICAL - Packet loss = 100%
[1363656444] SERVICE ALERT: HHHHH;Disk usage;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363656444] SERVICE ALERT: HHHHH;CPU Statistics;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363656444] SERVICE ALERT: HHHHH;Swap Usage;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363656485] SERVICE ALERT: HHHHH;Memory;CRITICAL;HARD;3;CRITICAL - 5.0% (828496 kB) free!
[1363656485] SERVICE ALERT: HHHHH;Current Users;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363656674] SERVICE ALERT: HHHHH;Memory;WARNING;HARD;3;WARNING - 6.0% (1974120 kB) free!
[1363656895] SERVICE ALERT: HHHHH;Total Processes;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363656984] SERVICE ALERT: HHHHH;PING;CRITICAL;HARD;1;PING CRITICAL - Packet loss = 100%
[1363657044] SERVICE ALERT: HHHHH;Disk usage;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363657044] SERVICE ALERT: HHHHH;CPU Statistics;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363657044] SERVICE ALERT: HHHHH;Swap Usage;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363657044] SERVICE ALERT: HHHHH;Total Processes;OK;HARD;3;PROCS OK: 409 processes
[1363657084] SERVICE ALERT: HHHHH;Current Users;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363657494] SERVICE ALERT: HHHHH;Total Processes;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363657584] SERVICE ALERT: HHHHH;PING;CRITICAL;HARD;1;PING CRITICAL - Packet loss = 100%
[1363657644] SERVICE ALERT: HHHHH;Disk usage;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363657644] SERVICE ALERT: HHHHH;CPU Statistics;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363657644] SERVICE ALERT: HHHHH;Swap Usage;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363657684] SERVICE ALERT: HHHHH;Current Users;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363658094] SERVICE ALERT: HHHHH;Total Processes;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363658184] SERVICE ALERT: HHHHH;PING;CRITICAL;HARD;1;PING CRITICAL - Packet loss = 100%
[1363658244] SERVICE ALERT: HHHHH;Disk usage;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363658244] SERVICE ALERT: HHHHH;CPU Statistics;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363658244] SERVICE ALERT: HHHHH;Swap Usage;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363658244] SERVICE ALERT: HHHHH;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 672 processes
[1363658285] SERVICE ALERT: HHHHH;Current Users;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363658364] SERVICE ALERT: HHHHH;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 709 processes
[1363658489] SERVICE ALERT: HHHHH;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 675 processes
[1363658695] SERVICE ALERT: HHHHH;Total Processes;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363658785] SERVICE ALERT: HHHHH;PING;CRITICAL;HARD;1;PING CRITICAL - Packet loss = 100%
[1363658844] SERVICE ALERT: HHHHH;Disk usage;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363658844] SERVICE ALERT: HHHHH;CPU Statistics;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363658844] SERVICE ALERT: HHHHH;Swap Usage;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363658885] SERVICE ALERT: HHHHH;Current Users;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363658995] SERVICE ALERT: HHHHH;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 603 processes
[1363659088] SERVICE ALERT: HHHHH;Total Processes;OK;HARD;3;PROCS OK: 456 processes
[1363659175] SERVICE ALERT: HHHHH;Memory;WARNING;HARD;3;WARNING - 7.1% (1171416 kB) free!
[1363659295] SERVICE ALERT: HHHHH;Total Processes;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363659384] SERVICE ALERT: HHHHH;PING;CRITICAL;HARD;1;PING CRITICAL - Packet loss = 100%
[1363659444] SERVICE ALERT: HHHHH;Disk usage;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363659444] SERVICE ALERT: HHHHH;CPU Statistics;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363659444] SERVICE ALERT: HHHHH;Swap Usage;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363659484] SERVICE ALERT: HHHHH;Current Users;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363659688] SERVICE ALERT: HHHHH;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 636 processes
[1363659805] SERVICE ALERT: HHHHH;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 690 processes
[1363659895] SERVICE ALERT: HHHHH;Total Processes;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363659924] SERVICE ALERT: HHHHH;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 654 processes
[1363659984] SERVICE ALERT: HHHHH;PING;CRITICAL;HARD;1;PING CRITICAL - Packet loss = 100%
[1363660044] SERVICE ALERT: HHHHH;Disk usage;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363660044] SERVICE ALERT: HHHHH;CPU Statistics;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363660044] SERVICE ALERT: HHHHH;Swap Usage;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363660085] SERVICE ALERT: HHHHH;Current Users;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363660274] SERVICE ALERT: HHHHH;Memory;OK;HARD;3;OK - 12.3% (4052920 kB) free.
[1363660494] SERVICE ALERT: HHHHH;Total Processes;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363660555] SERVICE ALERT: HHHHH;Disk usage;WARNING;HARD;3;DISK WARNING - free space: / 31976 MB (87% inode=99%): /home 3724 MB (39% inode=99%): /var 8909 MB (94% inode=99%): /usr 5810 MB (61% inode=93%): /u01 33760 MB (69% inode=97%): /archivelog 101039 MB (98% inode=99%): /data1 327036 MB (19% inode=99%): /redologs 49825 MB (97% inode=99%): /backup1 56361 MB (15% inode=99%): /index1 139337 MB (27% inode=99%): /data2 284843 MB (22% inode=99%):
[1363660584] SERVICE ALERT: HHHHH;PING;CRITICAL;HARD;1;PING CRITICAL - Packet loss = 100%
[1363660644] SERVICE ALERT: HHHHH;Disk usage;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363660644] SERVICE ALERT: HHHHH;CPU Statistics;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363660644] SERVICE ALERT: HHHHH;Swap Usage;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363660685] SERVICE ALERT: HHHHH;Current Users;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363660795] SERVICE ALERT: HHHHH;Total Processes;WARNING;HARD;3;PROCS WARNING: 599 processes
[1363661095] SERVICE ALERT: HHHHH;Total Processes;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363661185] SERVICE ALERT: HHHHH;PING;CRITICAL;HARD;1;PING CRITICAL - Packet loss = 100%
[1363661244] SERVICE ALERT: HHHHH;Disk usage;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363661244] SERVICE ALERT: HHHHH;CPU Statistics;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363661244] SERVICE ALERT: HHHHH;Swap Usage;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363661285] SERVICE ALERT: HHHHH;Current Users;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363661694] SERVICE ALERT: HHHHH;Total Processes;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363661784] SERVICE ALERT: HHHHH;PING;CRITICAL;HARD;1;PING CRITICAL - Packet loss = 100%
[1363661834] SERVICE ALERT: HHHHH;Current Load;WARNING;SOFT;1;WARNING - load average: 3.97, 5.13, 5.07
[1363661844] SERVICE ALERT: HHHHH;Disk usage;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363661844] SERVICE ALERT: HHHHH;CPU Statistics;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363661844] SERVICE ALERT: HHHHH;Swap Usage;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363661884] SERVICE ALERT: HHHHH;Current Users;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363661957] SERVICE ALERT: HHHHH;Current Load;OK;SOFT;2;OK - load average: 3.36, 4.52, 4.85
[1363661995] SERVICE ALERT: HHHHH;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 606 processes
[1363662294] SERVICE ALERT: HHHHH;Total Processes;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363662324] SERVICE ALERT: HHHHH;Total Processes;OK;HARD;3;PROCS OK: 398 processes
[1363662384] SERVICE ALERT: HHHHH;PING;CRITICAL;HARD;1;PING CRITICAL - Packet loss = 100%
[1363662444] SERVICE ALERT: HHHHH;Disk usage;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363662444] SERVICE ALERT: HHHHH;CPU Statistics;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363662444] SERVICE ALERT: HHHHH;Swap Usage;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363662485] SERVICE ALERT: HHHHH;Current Users;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363662895] SERVICE ALERT: HHHHH;Total Processes;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363662924] SERVICE ALERT: HHHHH;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 690 processes
[1363662984] SERVICE ALERT: HHHHH;PING;CRITICAL;HARD;1;PING CRITICAL - Packet loss = 100%
[1363663044] SERVICE ALERT: HHHHH;Disk usage;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363663044] SERVICE ALERT: HHHHH;CPU Statistics;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363663044] SERVICE ALERT: HHHHH;Swap Usage;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363663044] SERVICE ALERT: HHHHH;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 707 processes
[1363663085] SERVICE ALERT: HHHHH;Current Users;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363663164] SERVICE ALERT: HHHHH;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 701 processes
[1363663495] SERVICE ALERT: HHHHH;Total Processes;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363663585] SERVICE ALERT: HHHHH;PING;CRITICAL;HARD;1;CRITICAL - Plugin timed out after 10 seconds
[1363663644] SERVICE ALERT: HHHHH;Disk usage;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363663644] SERVICE ALERT: HHHHH;CPU Statistics;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363663644] SERVICE ALERT: HHHHH;Swap Usage;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363663685] SERVICE ALERT: HHHHH;Current Users;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363663694] SERVICE ALERT: HHHHH;Total Processes;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363663804] SERVICE ALERT: HHHHH;Total Processes;OK;SOFT;2;PROCS OK: 43 processes
[1363664094] SERVICE ALERT: HHHHH;Total Processes;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363664185] SERVICE ALERT: HHHHH;PING;CRITICAL;HARD;1;PING CRITICAL - Packet loss = 100%
[1363664244] SERVICE ALERT: HHHHH;Disk usage;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363664245] SERVICE ALERT: HHHHH;CPU Statistics;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363664245] SERVICE ALERT: HHHHH;Swap Usage;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363664284] SERVICE ALERT: HHHHH;Current Users;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363664395] SERVICE ALERT: HHHHH;Total Processes;WARNING;HARD;3;PROCS WARNING: 600 processes
[1363664694] SERVICE ALERT: HHHHH;Total Processes;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363664785] SERVICE ALERT: HHHHH;PING;CRITICAL;HARD;1;PING CRITICAL - Packet loss = 100%
[1363664844] SERVICE ALERT: HHHHH;Disk usage;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363664844] SERVICE ALERT: HHHHH;CPU Statistics;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363664844] SERVICE ALERT: HHHHH;Swap Usage;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363664884] SERVICE ALERT: HHHHH;Current Users;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363665295] SERVICE ALERT: HHHHH;Total Processes;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363665385] SERVICE ALERT: HHHHH;PING;CRITICAL;HARD;1;PING CRITICAL - Packet loss = 100%
[1363665445] SERVICE ALERT: HHHHH;Disk usage;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363665445] SERVICE ALERT: HHHHH;CPU Statistics;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363665445] SERVICE ALERT: HHHHH;Swap Usage;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363665484] SERVICE ALERT: HHHHH;Current Users;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363665714] SERVICE ALERT: HHHHH;Memory;WARNING;SOFT;1;WARNING - 8.5% (2800632 kB) free!
[1363665834] SERVICE ALERT: HHHHH;Memory;WARNING;SOFT;2;WARNING - 7.7% (2551016 kB) free!
[1363665894] SERVICE ALERT: HHHHH;Total Processes;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363665954] SERVICE ALERT: HHHHH;Memory;WARNING;HARD;3;WARNING - 8.0% (2636536 kB) free!
[1363665985] SERVICE ALERT: HHHHH;PING;CRITICAL;HARD;1;PING CRITICAL - Packet loss = 100%
[1363666004] SERVICE ALERT: HHHHH;Memory;WARNING;HARD;3;WARNING - 5.3% (871616 kB) free!
[1363666044] SERVICE ALERT: HHHHH;Disk usage;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363666044] SERVICE ALERT: HHHHH;CPU Statistics;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363666044] SERVICE ALERT: HHHHH;Swap Usage;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363666085] SERVICE ALERT: HHHHH;Current Users;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363666195] SERVICE ALERT: HHHHH;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 605 processes
[1363666495] SERVICE ALERT: HHHHH;Total Processes;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363666585] SERVICE ALERT: HHHHH;PING;CRITICAL;HARD;1;PING CRITICAL - Packet loss = 100%
[1363666615] SERVICE ALERT: HHHHH;Memory;OK;HARD;3;OK - 19.1% (794264 kB) free.
[1363666644] SERVICE ALERT: HHHHH;Disk usage;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363666644] SERVICE ALERT: HHHHH;CPU Statistics;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363666644] SERVICE ALERT: HHHHH;Swap Usage;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363666684] SERVICE ALERT: HHHHH;Current Users;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363667095] SERVICE ALERT: HHHHH;Total Processes;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363667185] SERVICE ALERT: HHHHH;PING;CRITICAL;HARD;1;PING CRITICAL - Packet loss = 100%
[1363667244] SERVICE ALERT: HHHHH;Disk usage;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363667244] SERVICE ALERT: HHHHH;CPU Statistics;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363667244] SERVICE ALERT: HHHHH;Swap Usage;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363667285] SERVICE ALERT: HHHHH;Current Users;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363667364] SERVICE ALERT: HHHHH;Total Processes;OK;HARD;3;PROCS OK: 461 processes
[1363667395] SERVICE ALERT: HHHHH;Total Processes;WARNING;HARD;3;PROCS WARNING: 600 processes
[1363667607] SERVICE ALERT: HHHHH;Memory;OK;HARD;3;OK - 11.5% (3789576 kB) free.
[1363667694] SERVICE ALERT: HHHHH;Total Processes;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363667745] SERVICE ALERT: HHHHH;Swap Usage;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363667785] SERVICE ALERT: HHHHH;PING;CRITICAL;HARD;1;PING CRITICAL - Packet loss = 100%
[1363667844] SERVICE ALERT: HHHHH;Disk usage;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363667844] SERVICE ALERT: HHHHH;CPU Statistics;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363667844] SERVICE ALERT: HHHHH;Swap Usage;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363667855] SERVICE ALERT: HHHHH;Swap Usage;OK;SOFT;2;SWAP OK - 100% free (1000 MB out of 1000 MB)
[1363667884] SERVICE ALERT: HHHHH;Current Users;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363667965] SERVICE ALERT: HHHHH;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 653 processes
[1363667995] SERVICE ALERT: HHHHH;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 603 processes
[1363668089] SERVICE ALERT: HHHHH;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 673 processes
[1363668205] SERVICE ALERT: HHHHH;Total Processes;OK;SOFT;3;PROCS OK: 407 processes
[1363668295] SERVICE ALERT: HHHHH;Total Processes;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363668384] SERVICE ALERT: HHHHH;PING;CRITICAL;HARD;1;PING CRITICAL - Packet loss = 100%
[1363668444] SERVICE ALERT: HHHHH;Disk usage;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363668444] SERVICE ALERT: HHHHH;CPU Statistics;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363668444] SERVICE ALERT: HHHHH;Swap Usage;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363668484] SERVICE ALERT: HHHHH;Current Users;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363668595] SERVICE ALERT: HHHHH;Total Processes;WARNING;HARD;3;PROCS WARNING: 600 processes
[1363668674] SERVICE FLAPPING ALERT: HHHHH;Memory;STOPPED; Service appears to have stopped flapping (4.4% change < 5.0% threshold)
[1363668804] SERVICE ALERT: HHHHH;Memory;WARNING;SOFT;1;WARNING - 10.0% (413272 kB) free!
[1363668805] SERVICE ALERT: HHHHH;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 661 processes
[1363668895] SERVICE ALERT: HHHHH;Total Processes;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363668925] SERVICE ALERT: HHHHH;Memory;WARNING;SOFT;2;WARNING - 10.0% (413256 kB) free!
[1363668925] SERVICE ALERT: HHHHH;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 695 processes
[1363668985] SERVICE ALERT: HHHHH;PING;CRITICAL;HARD;1;PING CRITICAL - Packet loss = 100%
[1363669044] SERVICE ALERT: HHHHH;Memory;WARNING;HARD;3;WARNING - 10.0% (413208 kB) free!
[1363669044] SERVICE ALERT: HHHHH;Disk usage;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363669044] SERVICE ALERT: HHHHH;CPU Statistics;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363669044] SERVICE ALERT: HHHHH;Swap Usage;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363669045] SERVICE ALERT: HHHHH;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 669 processes
[1363669084] SERVICE ALERT: HHHHH;Current Users;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363669195] SERVICE ALERT: HHHHH;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 605 processes
[1363669196] SERVICE ALERT: HHHHH;Total Processes;WARNING;SOFT;1;PROCS WARNING: 501 processes
[1363669314] SERVICE ALERT: HHHHH;Total Processes;WARNING;SOFT;2;PROCS WARNING: 502 processes
[1363669434] SERVICE ALERT: HHHHH;Total Processes;OK;SOFT;3;PROCS OK: 500 processes
[1363669495] SERVICE ALERT: HHHHH;Total Processes;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363669585] SERVICE ALERT: HHHHH;PING;CRITICAL;HARD;1;PING CRITICAL - Packet loss = 100%
[1363669604] SERVICE ALERT: HHHHH;Memory;CRITICAL;HARD;3;CRITICAL - 4.7% (775064 kB) free!
[1363669644] SERVICE ALERT: HHHHH;Disk usage;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363669644] SERVICE ALERT: HHHHH;CPU Statistics;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363669644] SERVICE ALERT: HHHHH;Swap Usage;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363669685] SERVICE ALERT: HHHHH;Current Users;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363670005] SERVICE ALERT: HHHHH;Memory;WARNING;SOFT;1;WARNING - 9.8% (3213368 kB) free!
[1363670035] SERVICE ALERT: HHHHH;Total Processes;WARNING;SOFT;1;PROCS WARNING: 503 processes
[1363670095] SERVICE ALERT: HHHHH;Total Processes;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363670124] SERVICE ALERT: HHHHH;Memory;WARNING;SOFT;2;WARNING - 9.8% (3204992 kB) free!
[1363670135] SERVICE ALERT: HHHHH;HTTP;WARNING;SOFT;1;HTTP WARNING: HTTP/1.1 200 OK - 10350 bytes in 5.992 second response time
[1363670154] SERVICE ALERT: HHHHH;Total Processes;WARNING;SOFT;2;PROCS WARNING: 506 processes
[1363670185] SERVICE ALERT: HHHHH;PING;CRITICAL;HARD;1;PING CRITICAL - Packet loss = 100%
[1363670244] SERVICE ALERT: HHHHH;Memory;WARNING;HARD;3;WARNING - 9.6% (3164632 kB) free!
[1363670244] SERVICE ALERT: HHHHH;Disk usage;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363670245] SERVICE ALERT: HHHHH;CPU Statistics;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363670245] SERVICE ALERT: HHHHH;Swap Usage;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363670254] SERVICE ALERT: HHHHH;HTTP;OK;SOFT;2;HTTP OK: HTTP/1.1 200 OK - 10350 bytes in 1.597 second response time
[1363670274] SERVICE ALERT: HHHHH;Total Processes;WARNING;HARD;3;PROCS WARNING: 505 processes
[1363670284] SERVICE ALERT: HHHHH;Current Users;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363670474] SERVICE ALERT: HHHHH;Memory;WARNING;SOFT;1;WARNING - 9.6% (3160392 kB) free!
[1363670595] SERVICE ALERT: HHHHH;Memory;WARNING;SOFT;2;WARNING - 8.3% (2721496 kB) free!
[1363670695] SERVICE ALERT: HHHHH;Total Processes;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363670714] SERVICE ALERT: HHHHH;Memory;WARNING;HARD;3;WARNING - 7.5% (2478648 kB) free!
[1363670724] SERVICE ALERT: HHHHH;HTTP;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363670785] SERVICE ALERT: HHHHH;PING;CRITICAL;HARD;1;PING CRITICAL - Packet loss = 100%
[1363670844] SERVICE ALERT: HHHHH;Memory;OK;HARD;3;OK - 10.2% (3341152 kB) free.
[1363670845] SERVICE ALERT: HHHHH;Disk usage;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363670845] SERVICE ALERT: HHHHH;CPU Statistics;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363670845] SERVICE ALERT: HHHHH;Swap Usage;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363670845] SERVICE ALERT: HHHHH;HTTP;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1363670845] SERVICE ALERT: HHHHH;Total Processes;OK;HARD;3;PROCS OK: 406 processes
[1363670885] SERVICE ALERT: HHHHH;Current Users;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363670964] SERVICE ALERT: HHHHH;HTTP;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1363671134] SERVICE ALERT: HHHHH;Current Load;WARNING;SOFT;1;WARNING - load average: 5.40, 6.13, 5.47
[1363671254] SERVICE ALERT: HHHHH;Current Load;WARNING;SOFT;2;WARNING - load average: 4.64, 5.58, 5.35
[1363671295] SERVICE ALERT: HHHHH;Total Processes;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363671375] SERVICE ALERT: HHHHH;Current Load;WARNING;HARD;3;WARNING - load average: 5.52, 5.51, 5.34
[1363671384] SERVICE ALERT: HHHHH;PING;CRITICAL;HARD;1;PING CRITICAL - Packet loss = 100%
[1363671444] SERVICE ALERT: HHHHH;Memory;WARNING;SOFT;1;WARNING - 10.0% (3272024 kB) free!
[1363671444] SERVICE ALERT: HHHHH;Disk usage;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363671444] SERVICE ALERT: HHHHH;CPU Statistics;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363671444] SERVICE ALERT: HHHHH;Swap Usage;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363671444] SERVICE ALERT: HHHHH;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 677 processes
[1363671484] SERVICE ALERT: HHHHH;Current Users;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363671555] SERVICE ALERT: HHHHH;Current Load;WARNING;SOFT;1;WARNING - load average: 7.38, 6.44, 5.58
[1363671559] SERVICE ALERT: HHHHH;HTTP;OK;HARD;3;HTTP OK: HTTP/1.1 302 Found - 470 bytes in 0.144 second response time
[1363671565] SERVICE ALERT: HHHHH;Memory;OK;SOFT;2;OK - 10.1% (3319824 kB) free.
[1363671565] SERVICE ALERT: HHHHH;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 652 processes
[1363671674] SERVICE ALERT: HHHHH;Current Load;WARNING;SOFT;2;WARNING - load average: 6.73, 6.70, 5.80
[1363671674] SERVICE ALERT: HHHHH;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 662 processes
[1363671689] SERVICE ALERT: HHHHH;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 688 processes
[1363671795] SERVICE ALERT: HHHHH;Current Load;WARNING;HARD;3;WARNING - load average: 6.12, 6.44, 5.81
[1363671796] SERVICE ALERT: HHHHH;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 657 processes
[1363671864] SERVICE ALERT: HHHHH;Current Users;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363671866] SERVICE ALERT: HHHHH;Memory;WARNING;HARD;3;WARNING - 5.7% (467392 kB) free!
[1363671894] SERVICE ALERT: HHHHH;Total Processes;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363671914] SERVICE ALERT: HHHHH;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 657 processes
[1363671954] SERVICE ALERT: HHHHH;Memory;CRITICAL;HARD;3;CRITICAL - 3.7% (1220352 kB) free!
[1363671975] SERVICE ALERT: HHHHH;Current Users;OK;SOFT;2;USERS OK - 1 users currently logged in
[1363671984] SERVICE ALERT: HHHHH;PING;CRITICAL;HARD;1;PING CRITICAL - Packet loss = 100%
[1363672044] SERVICE ALERT: HHHHH;Disk usage;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363672044] SERVICE ALERT: HHHHH;CPU Statistics;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363672044] SERVICE ALERT: HHHHH;Swap Usage;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363672084] SERVICE ALERT: HHHHH;Current Users;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363672134] HOST ALERT: HHHHH;DOWN;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363672164] HOST ALERT: HHHHH;UP;SOFT;2;NRPE v2.12
[1363672164] SERVICE ALERT: HHHHH;Memory;WARNING;SOFT;1;WARNING - 8.0% (2609680 kB) free!
[1363672290] SERVICE ALERT: HHHHH;Memory;WARNING;SOFT;2;WARNING - 7.5% (2447760 kB) free!
[1363672405] SERVICE ALERT: HHHHH;Memory;WARNING;HARD;3;WARNING - 7.2% (2374224 kB) free!
[1363672405] SERVICE FLAPPING ALERT: HHHHH;Memory;STARTED; Service appears to have started flapping (23.4% change >= 20.0% threshold)
[1363672454] SERVICE ALERT: HHHHH;Memory;CRITICAL;HARD;3;CRITICAL - 4.8% (3113232 kB) free!
[1363672454] SERVICE ALERT: HHHHH;Memory;CRITICAL;HARD;3;CRITICAL - 4.8% (3112472 kB) free!
[1363672455] SERVICE ALERT: HHHHH;Memory;CRITICAL;HARD;3;CRITICAL - 4.8% (3120808 kB) free!
[1363672495] SERVICE ALERT: HHHHH;Total Processes;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363672575] SERVICE ALERT: HHHHH;Current Load;OK;HARD;3;OK - load average: 3.50, 3.83, 4.49
[1363672584] SERVICE ALERT: HHHHH;Current Users;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363672585] SERVICE ALERT: HHHHH;PING;CRITICAL;HARD;1;PING CRITICAL - Packet loss = 100%
[1363672644] SERVICE ALERT: HHHHH;Disk usage;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363672644] SERVICE ALERT: HHHHH;CPU Statistics;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363672644] SERVICE ALERT: HHHHH;Swap Usage;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363672684] SERVICE ALERT: HHHHH;Current Users;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363672695] SERVICE ALERT: HHHHH;Current Users;OK;SOFT;2;USERS OK - 1 users currently logged in
[1363673095] SERVICE ALERT: HHHHH;Total Processes;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363673144] HOST ALERT: HHHHH;DOWN;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363673185] SERVICE ALERT: HHHHH;PING;CRITICAL;HARD;1;PING CRITICAL - Packet loss = 100%
[1363673214] HOST ALERT: HHHHH;UP;SOFT;2;NRPE v2.12
[1363673214] HOST FLAPPING ALERT: HHHHH;STARTED; Host appears to have started flapping (23.2% change > 20.0% threshold)
[1363673244] SERVICE ALERT: HHHHH;Disk usage;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363673244] SERVICE ALERT: HHHHH;CPU Statistics;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363673244] SERVICE ALERT: HHHHH;Swap Usage;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363673285] SERVICE ALERT: HHHHH;Current Users;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363673605] SERVICE ALERT: HHHHH;Memory;OK;HARD;3;OK - 11.2% (3680336 kB) free.
[1363673695] SERVICE ALERT: HHHHH;Total Processes;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363673714] SERVICE ALERT: HHHHH;Memory;OK;HARD;3;OK - 10.7% (3519752 kB) free.
[1363673785] SERVICE ALERT: HHHHH;PING;CRITICAL;HARD;1;PING CRITICAL - Packet loss = 100%
[1363673844] SERVICE ALERT: HHHHH;Disk usage;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363673844] SERVICE ALERT: HHHHH;CPU Statistics;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363673844] SERVICE ALERT: HHHHH;Swap Usage;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363673884] SERVICE ALERT: HHHHH;Current Users;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363674088] SERVICE ALERT: HHHHH;Total Processes;OK;HARD;3;PROCS OK: 441 processes
[1363674164] SERVICE ALERT: HHHHH;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 1415 processes
[1363674165] SERVICE ALERT: HHHHH;Memory;WARNING;SOFT;1;WARNING - 9.0% (2772088 kB) free!
[1363674205] SERVICE ALERT: HHHHH;Memory;WARNING;SOFT;1;WARNING - 9.4% (3098832 kB) free!
[1363674284] SERVICE ALERT: HHHHH;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 1415 processes
[1363674285] SERVICE ALERT: HHHHH;Memory;WARNING;SOFT;2;WARNING - 8.9% (2741936 kB) free!
[1363674295] SERVICE ALERT: HHHHH;Total Processes;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363674324] SERVICE ALERT: HHHHH;Memory;WARNING;SOFT;2;WARNING - 9.3% (3038048 kB) free!
[1363674375] SERVICE ALERT: HHHHH;Current Load;WARNING;SOFT;1;WARNING - load average: 4.42, 5.09, 5.05
[1363674375] SERVICE ALERT: HHHHH;PING;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 0.32 ms
[1363674384] HOST ALERT: HHHHH;UP;HARD;1;NRPE v2.12
[1363674404] SERVICE ALERT: HHHHH;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 1416 processes
[1363674404] SERVICE ALERT: HHHHH;Memory;WARNING;HARD;3;WARNING - 8.9% (2741544 kB) free!
[1363674434] SERVICE ALERT: HHHHH;Disk usage;OK;HARD;1;DISK OK - free space: / 2082 MB (52% inode=95%): /usr 1022 MB (17% inode=71%): /var 3771 MB (62% inode=96%): /movfiles 1076 MB (21% inode=99%): /oracle 6289 MB (20% inode=98%): /projectbkp 1032 MB (12% inode=97%): /oracle11g 2982 MB (14% inode=98%): /export/home 1227 MB (20% inode=97%):
[1363674434] SERVICE ALERT: HHHHH;Swap Usage;OK;HARD;1;SWAP OK - 100% free (8193 MB out of 8193 MB)
[1363674434] SERVICE ALERT: HHHHH;CPU Statistics;OK;HARD;1;CPU STATISTICS OK : user=0% system=0% iowait=0% idle=100%
[1363674444] SERVICE ALERT: HHHHH;Memory;OK;HARD;3;OK - 10.4% (427864 kB) free.
[1363674444] SERVICE ALERT: HHHHH;Memory;WARNING;HARD;3;WARNING - 9.4% (3079048 kB) free!
[1363674474] SERVICE ALERT: HHHHH;Current Users;OK;HARD;1;USERS OK - 0 users currently logged in
[1363674495] SERVICE ALERT: HHHHH;Current Load;OK;SOFT;2;OK - load average: 4.58, 4.87, 4.97
[1363674527] SERVICE ALERT: HHHHH;Memory;OK;HARD;3;OK - 88.4% (3629008 kB) free.
[1363674688] SERVICE ALERT: HHHHH;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 685 processes
[1363674714] SERVICE ALERT: HHHHH;PING;WARNING;SOFT;1;PING WARNING - Packet loss = 20%, RTA = 5.60 ms
[1363674806] SERVICE ALERT: HHHHH;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 702 processes
[1363674834] SERVICE ALERT: HHHHH;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 3.87 ms
[1363674854] SERVICE ALERT: HHHHH;Memory;WARNING;HARD;3;WARNING - 6.5% (4243808 kB) free!
[1363674854] SERVICE ALERT: HHHHH;Memory;WARNING;HARD;3;WARNING - 6.5% (4242880 kB) free!
[1363674855] SERVICE ALERT: HHHHH;Memory;WARNING;HARD;3;WARNING - 6.5% (4248872 kB) free!
[1363674866] SERVICE ALERT: HHHHH;Memory;CRITICAL;HARD;3;CRITICAL - 3.4% (279776 kB) free!
[1363674885] SERVICE ALERT: HHHHH;Current Load;OK;HARD;3;OK - load average: 0.03, 0.09, 0.08
[1363674885] SERVICE ALERT: HHHHH;Total Processes;OK;HARD;1;PROCS OK: 53 processes
[1363674914] SERVICE ALERT: HHHHH;Memory;WARNING;SOFT;1;WARNING - 7.1% (2324520 kB) free!
[1363674924] SERVICE ALERT: HHHHH;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 674 processes
[1363675034] SERVICE ALERT: HHHHH;Memory;WARNING;SOFT;2;WARNING - 9.6% (3161248 kB) free!
[1363675044] SERVICE ALERT: HHHHH;Memory;OK;HARD;3;OK - 10.2% (3346464 kB) free.
[1363675095] SERVICE ALERT: HHHHH;Current Load;WARNING;SOFT;1;WARNING - load average: 4.69, 5.43, 5.25
[1363675158] SERVICE ALERT: HHHHH;Memory;WARNING;HARD;3;WARNING - 6.5% (2129056 kB) free!
[1363675214] SERVICE ALERT: HHHHH;Current Load;WARNING;SOFT;2;WARNING - load average: 4.30, 5.03, 5.13
[1363675334] SERVICE ALERT: HHHHH;Current Load;WARNING;HARD;3;WARNING - load average: 4.57, 4.85, 5.04
[1363675466] SERVICE ALERT: HHHHH;Memory;WARNING;HARD;3;WARNING - 5.9% (489088 kB) free!
[1363675564] SERVICE ALERT: HHHHH;PING;WARNING;SOFT;1;PING WARNING - Packet loss = 0%, RTA = 277.00 ms
[1363675604] SERVICE ALERT: HHHHH;Memory;OK;HARD;3;OK - 22.9% (7069472 kB) free.
[1363675644] SERVICE ALERT: HHHHH;PING;WARNING;SOFT;1;PING WARNING - Packet loss = 20%, RTA = 37.60 ms
[1363675644] SERVICE ALERT: HHHHH;Memory;WARNING;SOFT;1;WARNING - 6.9% (2272152 kB) free!
[1363675685] SERVICE ALERT: HHHHH;PING;WARNING;SOFT;2;PING WARNING - Packet loss = 0%, RTA = 204.00 ms
[1363675764] SERVICE ALERT: HHHHH;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 37.50 ms
[1363675764] SERVICE ALERT: HHHHH;Memory;WARNING;SOFT;2;WARNING - 6.1% (2001640 kB) free!
[1363675814] SERVICE ALERT: HHHHH;PING;WARNING;HARD;3;PING WARNING - Packet loss = 0%, RTA = 284.00 ms
[1363675889] SERVICE ALERT: HHHHH;Memory;WARNING;HARD;3;WARNING - 5.2% (1692648 kB) free!
[1363675914] SERVICE ALERT: HHHHH;PING;WARNING;SOFT;1;PING WARNING - Packet loss = 0%, RTA = 251.00 ms
[1363675934] SERVICE ALERT: HHHHH;Current Load;OK;HARD;3;OK - load average: 3.14, 3.46, 4.22
[1363675934] SERVICE FLAPPING ALERT: HHHHH;Current Load;STARTED; Service appears to have started flapping (23.3% change >= 20.0% threshold)
[1363676034] SERVICE ALERT: HHHHH;PING;WARNING;SOFT;2;PING WARNING - Packet loss = 0%, RTA = 278.00 ms
[1363676125] SERVICE ALERT: HHHHH;Total Processes;OK;HARD;3;PROCS OK: 415 processes
[1363676154] SERVICE ALERT: HHHHH;PING;WARNING;HARD;3;PING WARNING - Packet loss = 0%, RTA = 271.00 ms
[1363676204] SERVICE ALERT: HHHHH;Memory;WARNING;SOFT;1;WARNING - 8.7% (2684080 kB) free!
[1363676326] SERVICE ALERT: HHHHH;Memory;WARNING;SOFT;2;WARNING - 8.7% (2699176 kB) free!
[1363676414] SERVICE ALERT: HHHHH;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 18.80 ms
[1363676444] SERVICE ALERT: HHHHH;Memory;WARNING;HARD;3;WARNING - 9.9% (3063216 kB) free!
[1363676724] SERVICE ALERT: HHHHH;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 681 processes
[1363676755] SERVICE ALERT: HHHHH;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 13.70 ms
[1363676814] SERVICE ALERT: HHHHH;Memory;WARNING;SOFT;1;WARNING - 9.8% (407768 kB) free!
[1363676844] SERVICE ALERT: HHHHH;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 677 processes
[1363676934] SERVICE ALERT: HHHHH;Memory;WARNING;SOFT;2;WARNING - 9.9% (410576 kB) free!
[1363676955] SERVICE ALERT: HHHHH;Total Processes;OK;HARD;3;PROCS OK: 102 processes
[1363676964] SERVICE ALERT: HHHHH;Memory;WARNING;HARD;3;WARNING - 6.1% (496600 kB) free!
[1363676965] SERVICE ALERT: HHHHH;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 677 processes
[1363677006] SERVICE ALERT: HHHHH;Total Processes;OK;HARD;3;PROCS OK: 78 processes
[1363677054] SERVICE ALERT: HHHHH;Memory;OK;SOFT;3;OK - 10.1% (419248 kB) free.
[1363677089] SERVICE ALERT: HHHHH;Memory;OK;HARD;3;OK - 11.0% (3608056 kB) free.
[1363677254] SERVICE ALERT: HHHHH;Memory;CRITICAL;HARD;3;CRITICAL - 4.3% (2837144 kB) free!
[1363677254] SERVICE ALERT: HHHHH;Memory;CRITICAL;HARD;3;CRITICAL - 4.3% (2837568 kB) free!
[1363677255] SERVICE ALERT: HHHHH;Memory;CRITICAL;HARD;3;CRITICAL - 4.3% (2840632 kB) free!
[1363677404] SERVICE ALERT: HHHHH;Memory;WARNING;HARD;3;WARNING - 7.5% (1234568 kB) free!
[1363677564] SERVICE ALERT: HHHHH;Memory;CRITICAL;HARD;3;CRITICAL - 3.2% (261624 kB) free!
[1363677644] SERVICE ALERT: HHHHH;Memory;OK;HARD;3;OK - 17.1% (5267936 kB) free.
[1363677644] SERVICE FLAPPING ALERT: HHHHH;Memory;STARTED; Service appears to have started flapping (23.7% change >= 20.0% threshold)
[1363677654] SERVICE ALERT: HHHHH;Memory;WARNING;SOFT;1;WARNING - 9.5% (396768 kB) free!
[1363677688] SERVICE ALERT: HHHHH;Memory;WARNING;SOFT;1;WARNING - 10.0% (3279848 kB) free!
[1363677775] SERVICE ALERT: HHHHH;Memory;WARNING;SOFT;2;WARNING - 9.5% (395280 kB) free!
[1363677805] SERVICE ALERT: HHHHH;Memory;WARNING;SOFT;2;WARNING - 9.1% (2998512 kB) free!
[1363677895] SERVICE ALERT: HHHHH;Memory;WARNING;HARD;3;WARNING - 9.8% (408544 kB) free!
[1363677924] SERVICE ALERT: HHHHH;Memory;WARNING;HARD;3;WARNING - 9.9% (3244136 kB) free!
[1363677934] SERVICE ALERT: HHHHH;Memory;OK;HARD;3;OK - 35.4% (11606640 kB) free.
[1363678164] SERVICE ALERT: HHHHH;Memory;WARNING;HARD;3;WARNING - 5.8% (475424 kB) free!
[1363678244] SERVICE ALERT: HHHHH;Memory;WARNING;SOFT;1;WARNING - 9.2% (2845896 kB) free!
[1363678365] SERVICE ALERT: HHHHH;Memory;WARNING;SOFT;2;WARNING - 8.9% (2760328 kB) free!
[1363678375] SERVICE ALERT: HHHHH;Current Load;WARNING;SOFT;1;WARNING - load average: 15.10, 6.09, 3.34
[1363678464] SERVICE ALERT: HHHHH;Disk usage;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363678485] SERVICE ALERT: HHHHH;Memory;WARNING;HARD;3;WARNING - 8.9% (2757592 kB) free!
[1363678495] SERVICE ALERT: HHHHH;Current Load;OK;SOFT;2;OK - load average: 11.12, 7.80, 4.34
[1363678524] SERVICE ALERT: HHHHH;Memory;CRITICAL;HARD;3;CRITICAL - 3.7% (1201216 kB) free!
[1363678534] SERVICE ALERT: HHHHH;Memory;CRITICAL;SOFT;1;CRITICAL - 5.0% (1629600 kB) free!
[1363678575] SERVICE ALERT: HHHHH;Disk usage;OK;SOFT;2;DISK OK - free space: / 421 MB (8% inode=73%):
[1363678634] SERVICE ALERT: HHHHH;PING;WARNING;SOFT;1;PING WARNING - Packet loss = 20%, RTA = 4.38 ms
[1363678655] SERVICE ALERT: HHHHH;Memory;CRITICAL;SOFT;2;CRITICAL - 3.9% (1276584 kB) free!
[1363678666] SERVICE ALERT: HHHHH;Swap Usage;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363678674] HOST FLAPPING ALERT: HHHHH;STOPPED; Host appears to have stopped flapping (3.8% change < 5.0% threshold)
[1363678746] SERVICE ALERT: HHHHH;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 3.62 ms
[1363678764] SERVICE ALERT: HHHHH;Current Load;WARNING;SOFT;1;WARNING - load average: 4.85, 11.87, 9.00
[1363678765] SERVICE ALERT: HHHHH;Memory;OK;HARD;3;OK - 67.0% (5462488 kB) free.
[1363678765] SERVICE FLAPPING ALERT: HHHHH;Memory;STARTED; Service appears to have started flapping (24.2% change >= 20.0% threshold)
[1363678778] SERVICE ALERT: HHHHH;Memory;CRITICAL;HARD;3;CRITICAL - 3.6% (1191784 kB) free!
[1363678780] SERVICE ALERT: HHHHH;Swap Usage;OK;SOFT;2;SWAP OK - 100% free (1000 MB out of 1000 MB)
[1363678889] SERVICE ALERT: HHHHH;Current Load;WARNING;SOFT;2;WARNING - load average: 4.00, 9.27, 8.37
[1363679005] SERVICE ALERT: HHHHH;Current Load;WARNING;HARD;3;WARNING - load average: 3.25, 7.32, 7.75
[1363679124] SERVICE ALERT: HHHHH;Memory;WARNING;HARD;3;WARNING - 9.2% (3022768 kB) free!
[1363679294] SERVICE ALERT: HHHHH;PING;WARNING;SOFT;1;PING WARNING - Packet loss = 44%, RTA = 55.55 ms
[1363679345] SERVICE ALERT: HHHHH;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 55.69 ms
[1363679364] SERVICE ALERT: HHHHH;Current Load;CRITICAL;SOFT;1;CRITICAL - load average: 63.76, 27.11, 11.45
[1363679489] SERVICE ALERT: HHHHH;Current Load;CRITICAL;SOFT;2;CRITICAL - load average: 71.39, 42.29, 19.19
[1363679524] SERVICE ALERT: HHHHH;Swap Usage;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363679544] HOST ALERT: HHHHH;DOWN;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363679604] HOST ALERT: HHHHH;UP;SOFT;2;NRPE v2.10
[1363679605] SERVICE ALERT: HHHHH;Current Load;CRITICAL;HARD;3;CRITICAL - load average: 29.02, 35.46, 19.56
[1363679634] SERVICE ALERT: HHHHH;Swap Usage;OK;SOFT;2;SWAP OK - 53% free (25426 MB out of 48137 MB)  [20077 (62%)] [5349 (35%)]
[1363679666] SERVICE ALERT: HHHHH;Memory;CRITICAL;HARD;3;CRITICAL - 4.2% (348600 kB) free!
[1363679666] SERVICE FLAPPING ALERT: HHHHH;Memory;STARTED; Service appears to have started flapping (21.3% change >= 20.0% threshold)
[1363679964] SERVICE ALERT: HHHHH;Memory;CRITICAL;SOFT;1;CRITICAL - 2.4% (193176 kB) free!
[1363680089] SERVICE ALERT: HHHHH;Memory;CRITICAL;SOFT;2;CRITICAL - 2.1% (168080 kB) free!
[1363680206] SERVICE ALERT: HHHHH;Current Load;OK;HARD;3;OK - load average: 2.66, 2.83, 4.18
[1363680207] SERVICE ALERT: HHHHH;Memory;CRITICAL;HARD;3;CRITICAL - 1.5% (125072 kB) free!
[1363680324] SERVICE ALERT: HHHHH;Memory;CRITICAL;HARD;3;CRITICAL - 1.9% (631368 kB) free!
[1363680558] SERVICE ALERT: HHHHH;Memory;CRITICAL;HARD;3;CRITICAL - 3.2% (1049352 kB) free!
[1363680558] SERVICE FLAPPING ALERT: HHHHH;Memory;STARTED; Service appears to have started flapping (20.3% change >= 20.0% threshold)
[1363680595] SERVICE FLAPPING ALERT: HHHHH;Total Processes;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1363680805] SERVICE ALERT: HHHHH;Current Load;WARNING;HARD;3;WARNING - load average: 1.47, 5.82, 13.95
[1363681485] SERVICE ALERT: HHHHH;Memory;OK;HARD;3;OK - 19.1% (5911784 kB) free.
[1363682006] SERVICE ALERT: HHHHH;Current Load;OK;HARD;3;OK - load average: 0.44, 0.40, 3.96
[1363682685] SERVICE ALERT: HHHHH;Memory;WARNING;SOFT;1;WARNING - 9.4% (2895672 kB) free!
[1363682804] SERVICE ALERT: HHHHH;Memory;WARNING;SOFT;2;WARNING - 9.4% (2895664 kB) free!
[1363682926] SERVICE ALERT: HHHHH;Memory;WARNING;HARD;3;WARNING - 9.4% (2893368 kB) free!
[1363682958] SERVICE ALERT: HHHHH;Memory;WARNING;HARD;3;WARNING - 8.7% (2866472 kB) free!
[1363684544] SERVICE ALERT: HHHHH;Current Load;WARNING;SOFT;1;WARNING - load average: 5.38, 5.43, 5.21
[1363684666] SERVICE ALERT: HHHHH;Current Load;WARNING;SOFT;2;WARNING - load average: 5.42, 5.40, 5.23
[1363684726] SERVICE ALERT: HHHHH;Memory;OK;HARD;3;OK - 10.3% (3188528 kB) free.
[1363684758] SERVICE ALERT: HHHHH;Disk usage;UNKNOWN;SOFT;1;CHECK_NRPE: Received 0 bytes from daemon.  Check the remote server logs for error messages.
[1363684758] SERVICE ALERT: HHHHH;Total Processes;UNKNOWN;SOFT;1;CHECK_NRPE: Received 0 bytes from daemon.  Check the remote server logs for error messages.
[1363684758] SERVICE ALERT: HHHHH;Current Load;WARNING;SOFT;1;WARNING - load average: 7.56, 7.28, 5.45
[1363684784] SERVICE ALERT: HHHHH;Current Load;WARNING;HARD;3;WARNING - load average: 5.52, 5.44, 5.27
[1363684874] SERVICE ALERT: HHHHH;Total Processes;OK;SOFT;2;PROCS OK: 49 processes
[1363684874] SERVICE ALERT: HHHHH;Disk usage;OK;SOFT;2;DISK OK - free space: / 95047 MB (94% inode=99%): /export 95047 MB (99% inode=99%): /export/home 95047 MB (99% inode=99%): /opt 95047 MB (77% inode=99%): /rpool 95047 MB (99% inode=99%):
[1363684877] SERVICE ALERT: HHHHH;Current Load;WARNING;SOFT;2;WARNING - load average: 7.46, 7.32, 5.69
[1363684995] SERVICE ALERT: HHHHH;Current Load;WARNING;HARD;3;WARNING - load average: 8.49, 7.78, 6.07
[1363685124] SERVICE ALERT: HHHHH;Memory;WARNING;HARD;3;WARNING - 5.5% (1808232 kB) free!
[1363685328] SERVICE ALERT: HHHHH;Memory;WARNING;SOFT;1;WARNING - 8.9% (2756696 kB) free!
[1363685444] SERVICE ALERT: HHHHH;Memory;WARNING;SOFT;2;WARNING - 8.6% (2662216 kB) free!
[1363685565] SERVICE ALERT: HHHHH;Memory;WARNING;HARD;3;WARNING - 8.7% (2699880 kB) free!
[1363685784] SERVICE ALERT: HHHHH;Disk usage;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363685896] SERVICE ALERT: HHHHH;Disk usage;OK;SOFT;2;DISK OK - free space: / 421 MB (8% inode=73%):
[1363685984] SERVICE ALERT: HHHHH;Current Load;OK;HARD;3;OK - load average: 4.57, 4.64, 4.84
[1363686324] SERVICE ALERT: HHHHH;Memory;CRITICAL;HARD;3;CRITICAL - 4.3% (1409256 kB) free!
[1363686584] SERVICE ALERT: HHHHH;Total Processes;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363686584] SERVICE ALERT: HHHHH;Current Load;WARNING;SOFT;1;WARNING - load average: 5.36, 5.19, 5.01
[1363686604] HOST ALERT: HHHHH;DOWN;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363686677] HOST ALERT: HHHHH;DOWN;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1363686694] SERVICE ALERT: HHHHH;CPU Statistics;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363686704] SERVICE ALERT: HHHHH;Total Processes;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1363686704] SERVICE ALERT: HHHHH;Current Load;WARNING;SOFT;2;WARNING - load average: 5.22, 5.21, 5.04
[1363686744] HOST ALERT: HHHHH;UP;SOFT;3;NRPE v2.10
[1363686805] SERVICE ALERT: HHHHH;Current Load;CRITICAL;SOFT;1;CRITICAL - load average: 43.36, 32.03, 19.22
[1363686805] SERVICE ALERT: HHHHH;Memory;WARNING;HARD;3;WARNING - 6.6% (541088 kB) free!
[1363686824] SERVICE ALERT: HHHHH;Current Load;WARNING;HARD;3;WARNING - load average: 5.66, 5.35, 5.11
[1363686925] SERVICE ALERT: HHHHH;Current Load;CRITICAL;SOFT;2;CRITICAL - load average: 35.65, 33.69, 21.50
[1363686964] SERVICE ALERT: HHHHH;Memory;WARNING;HARD;3;WARNING - 19.9% (6543640 kB) free!
[1363687045] SERVICE ALERT: HHHHH;Current Load;CRITICAL;HARD;3;CRITICAL - load average: 32.84, 33.36, 22.91
[1363687045] SERVICE FLAPPING ALERT: HHHHH;Current Load;STARTED; Service appears to have started flapping (21.1% change >= 20.0% threshold)
[1363687294] SERVICE ALERT: HHHHH;CPU Statistics;OK;HARD;1;CPU STATISTICS OK : user=73% system=11% iowait=15% idle=2%
[1363687304] SERVICE ALERT: HHHHH;Total Processes;OK;HARD;3;PROCS OK: 113 processes
[1363687334] SERVICE FLAPPING ALERT: HHHHH;Current Load;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1363687407] SERVICE ALERT: HHHHH;Memory;CRITICAL;HARD;3;CRITICAL - 1.5% (125280 kB) free!
[1363687425] SERVICE ALERT: HHHHH;Current Load;OK;HARD;3;OK - load average: 4.56, 4.75, 4.93
[1363687425] SERVICE FLAPPING ALERT: HHHHH;Current Load;STARTED; Service appears to have started flapping (24.1% change >= 20.0% threshold)
[1363687466] SERVICE FLAPPING ALERT: HHHHH;Memory;STOPPED; Service appears to have stopped flapping (4.5% change < 5.0% threshold)
[1363687524] SERVICE ALERT: HHHHH;Memory;WARNING;HARD;3;WARNING - 5.3% (1734232 kB) free!
[1363687564] SERVICE ALERT: HHHHH;Disk usage;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363687644] SERVICE ALERT: HHHHH;Current Load;WARNING;HARD;3;WARNING - load average: 2.77, 9.87, 16.25
[1363687674] SERVICE ALERT: HHHHH;Disk usage;OK;SOFT;2;DISK OK - free space: / 32038 MB (87% inode=99%): /home 7533 MB (80% inode=99%): /usr 5810 MB (61% inode=93%): /var 8899 MB (94% inode=99%): /u01 31006 MB (63% inode=98%): /archivelog 102127 MB (99% inode=99%): /data1 767377 MB (74% inode=99%): /redologs 45253 MB (88% inode=99%): /backup1 75570 MB (36% inode=99%): /index1 242359 MB (94% inode=99%): /data2 223976 MB (87% inode=99%):
[1363688124] SERVICE ALERT: HHHHH;Memory;CRITICAL;HARD;3;CRITICAL - 3.6% (1183456 kB) free!
[1363688145] SERVICE ALERT: HHHHH;CPU Statistics;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363688164] SERVICE ALERT: HHHHH;Memory;CRITICAL;HARD;3;CRITICAL - 11.0% (3618864 kB) free!
[1363688204] SERVICE ALERT: HHHHH;Total Processes;OK;HARD;3;PROCS OK: 159 processes
[1363688264] SERVICE ALERT: HHHHH;CPU Statistics;OK;SOFT;2;CPU STATISTICS OK : user=0% system=2% iowait=0% idle=98%
[1363688364] SERVICE FLAPPING ALERT: HHHHH;Total Processes;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1363688714] SERVICE ALERT: HHHHH;Total Processes;OK;HARD;3;PROCS OK: 240 processes
[1363688724] SERVICE ALERT: HHHHH;Memory;OK;HARD;3;OK - 11.0% (3610624 kB) free.
[1363688844] SERVICE ALERT: HHHHH;Current Load;OK;HARD;3;OK - load average: 0.19, 0.64, 4.94
[1363688976] SERVICE ALERT: HHHHH;Disk usage;WARNING;HARD;3;DISK WARNING - free space: / 10078 MB (53% inode=97%): /u01 10636 MB (8% inode=99%):
[1363689254] SERVICE ALERT: HHHHH;Memory;WARNING;HARD;3;WARNING - 5.1% (3329696 kB) free!
[1363689254] SERVICE ALERT: HHHHH;Memory;WARNING;HARD;3;WARNING - 5.1% (3333080 kB) free!
[1363689255] SERVICE ALERT: HHHHH;Memory;WARNING;HARD;3;WARNING - 5.1% (3335776 kB) free!
[1363689504] SERVICE ALERT: HHHHH;Current Users;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363689504] SERVICE ALERT: HHHHH;Disk usage;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363689615] SERVICE ALERT: HHHHH;Disk usage;OK;SOFT;2;DISK OK - free space: / 421 MB (8% inode=73%):
[1363689615] SERVICE ALERT: HHHHH;Current Users;OK;SOFT;2;USERS OK - 1 users currently logged in
[1363689774] SERVICE ALERT: HHHHH;Memory;CRITICAL;HARD;3;CRITICAL - 4.6% (764648 kB) free!
[1363689854] SERVICE ALERT: HHHHH;Memory;CRITICAL;HARD;3;CRITICAL - 4.2% (2768632 kB) free!
[1363689854] SERVICE ALERT: HHHHH;Memory;CRITICAL;HARD;3;CRITICAL - 4.2% (2767312 kB) free!
[1363689855] SERVICE ALERT: HHHHH;Memory;CRITICAL;HARD;3;CRITICAL - 4.2% (2775040 kB) free!
[1363689914] SERVICE ALERT: HHHHH;Total Processes;WARNING;SOFT;1;PROCS WARNING: 569 processes
[1363690034] SERVICE ALERT: HHHHH;Total Processes;WARNING;SOFT;2;PROCS WARNING: 569 processes
[1363690159] SERVICE ALERT: HHHHH;Total Processes;WARNING;HARD;3;PROCS WARNING: 569 processes
[1363690374] SERVICE ALERT: HHHHH;Memory;WARNING;HARD;3;WARNING - 7.3% (1196536 kB) free!
[1363690554] SERVICE ALERT: HHHHH;PING;WARNING;SOFT;1;PING WARNING - Packet loss = 20%, RTA = 8.15 ms
[1363690675] SERVICE ALERT: HHHHH;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 19.50 ms
[1363690965] SERVICE ALERT: HHHHH;Memory;OK;HARD;3;OK - 21.9% (6752016 kB) free.
[1363691204] HOST ALERT: HHHHH;DOWN;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363691274] HOST ALERT: HHHHH;UP;SOFT;2;NRPE v2.12
[1363691358] SERVICE ALERT: HHHHH;Memory;CRITICAL;HARD;3;CRITICAL - 4.7% (1539656 kB) free!
[1363691764] SERVICE ALERT: HHHHH;Total Processes;WARNING;SOFT;1;PROCS WARNING: 502 processes
[1363691884] SERVICE ALERT: HHHHH;Total Processes;WARNING;SOFT;2;PROCS WARNING: 509 processes
[1363691958] SERVICE ALERT: HHHHH;Memory;WARNING;HARD;3;WARNING - 6.5% (2135152 kB) free!
[1363692004] SERVICE ALERT: HHHHH;Total Processes;WARNING;HARD;3;PROCS WARNING: 510 processes
[1363692559] SERVICE ALERT: HHHHH;Memory;CRITICAL;HARD;3;CRITICAL - 4.3% (1432344 kB) free!
[1363692604] SERVICE ALERT: HHHHH;Total Processes;OK;HARD;3;PROCS OK: 500 processes
[1363692686] SERVICE ALERT: HHHHH;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 603 processes
[1363693144] SERVICE ALERT: HHHHH;Disk usage;WARNING;SOFT;1;DISK WARNING - free space: / 38901 MB (82% inode=98%): /var 8922 MB (94% inode=99%): /home 4789 MB (50% inode=99%): /u01 25400 MB (53% inode=98%): /data1 374677 MB (36% inode=99%): /backup1 74545 MB (29% inode=99%): /u02 35645 MB (58% inode=99%): /data2 573804 MB (56% inode=99%): /index1 214373 MB (20% inode=99%): /archivelog 97638 MB (95% inode=99%): /redologs 46820 MB (91% inode=99%):
[1363693164] SERVICE ALERT: HHHHH;Total Processes;OK;HARD;3;PROCS OK: 410 processes
[1363693204] SERVICE ALERT: HHHHH;PING;WARNING;SOFT;1;PING WARNING - Packet loss = 37%, RTA = 8.36 ms
[1363693204] SERVICE ALERT: HHHHH;Total Processes;WARNING;SOFT;1;PROCS WARNING: 503 processes
[1363693264] SERVICE ALERT: HHHHH;Disk usage;WARNING;SOFT;2;DISK WARNING - free space: / 38901 MB (82% inode=98%): /var 8922 MB (94% inode=99%): /home 4789 MB (50% inode=99%): /u01 25400 MB (53% inode=98%): /data1 374677 MB (36% inode=99%): /backup1 74545 MB (29% inode=99%): /u02 35645 MB (58% inode=99%): /data2 566007 MB (55% inode=99%): /index1 206399 MB (20% inode=99%): /archivelog 97638 MB (95% inode=99%): /redologs 46820 MB (91% inode=99%):
[1363693287] SERVICE ALERT: HHHHH;Total Processes;WARNING;HARD;3;PROCS WARNING: 590 processes
[1363693314] SERVICE ALERT: HHHHH;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 3.50 ms
[1363693326] SERVICE ALERT: HHHHH;Total Processes;OK;SOFT;2;PROCS OK: 500 processes
[1363693384] SERVICE ALERT: HHHHH;Disk usage;WARNING;HARD;3;DISK WARNING - free space: / 38901 MB (82% inode=98%): /var 8922 MB (94% inode=99%): /home 4789 MB (50% inode=99%): /u01 25400 MB (53% inode=98%): /data1 374677 MB (36% inode=99%): /backup1 74545 MB (29% inode=99%): /u02 35645 MB (58% inode=99%): /data2 558860 MB (54% inode=99%): /index1 199061 MB (19% inode=99%): /archivelog 97638 MB (95% inode=99%): /redologs 46820 MB (91% inode=99%):
[1363693425] SERVICE ALERT: HHHHH;Current Load;WARNING;SOFT;1;WARNING - load average: 5.27, 5.37, 5.08
[1363693545] SERVICE ALERT: HHHHH;Current Load;WARNING;SOFT;2;WARNING - load average: 5.36, 5.38, 5.12
[1363693667] SERVICE ALERT: HHHHH;Current Load;WARNING;HARD;3;WARNING - load average: 4.84, 5.17, 5.07
[1363693758] SERVICE ALERT: HHHHH;Memory;WARNING;HARD;3;WARNING - 6.0% (1966016 kB) free!
[1363693764] SERVICE ALERT: HHHHH;Current Load;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363693764] SERVICE ALERT: HHHHH;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 724 processes
[1363693875] SERVICE ALERT: HHHHH;Current Load;OK;SOFT;2;OK - load average: 0.01, 0.01, 0.00
[1363693875] SERVICE ALERT: HHHHH;Disk usage;OK;HARD;3;DISK OK - free space: / 15676 MB (31% inode=95%): /app01 5692 MB (11% inode=91%): /app02 15855 MB (31% inode=98%):
[1363693890] SERVICE ALERT: HHHHH;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 638 processes
[1363693894] SERVICE ALERT: HHHHH;CPU Statistics;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363694005] SERVICE ALERT: HHHHH;Memory;OK;HARD;3;OK - 68.0% (5548672 kB) free.
[1363694005] SERVICE ALERT: HHHHH;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 704 processes
[1363694006] SERVICE ALERT: HHHHH;CPU Statistics;OK;SOFT;2;CPU STATISTICS OK : user=25% system=12% iowait=18% idle=46%
[1363694267] SERVICE ALERT: HHHHH;Current Load;OK;HARD;3;OK - load average: 4.76, 4.90, 4.99
[1363694816] HOST ALERT: HHHHH;DOWN;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363694854] SERVICE ALERT: HHHHH;Current Load;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363694894] HOST ALERT: HHHHH;UP;SOFT;2;NRPE v2.10
[1363695126] SERVICE ALERT: HHHHH;Total Processes;WARNING;SOFT;1;PROCS WARNING: 508 processes
[1363695205] SERVICE ALERT: HHHHH;Memory;CRITICAL;SOFT;1;CRITICAL - 1.7% (141024 kB) free!
[1363695245] SERVICE ALERT: HHHHH;Total Processes;WARNING;SOFT;2;PROCS WARNING: 509 processes
[1363695324] SERVICE ALERT: HHHHH;Memory;CRITICAL;SOFT;2;CRITICAL - 3.4% (278744 kB) free!
[1363695365] SERVICE ALERT: HHHHH;Total Processes;WARNING;HARD;3;PROCS WARNING: 508 processes
[1363695444] SERVICE ALERT: HHHHH;Current Load;CRITICAL;SOFT;1;CRITICAL - load average: 13.82, 27.82, 24.26
[1363695444] SERVICE ALERT: HHHHH;Memory;CRITICAL;HARD;3;CRITICAL - 1.5% (124584 kB) free!
[1363695558] SERVICE ALERT: HHHHH;Memory;CRITICAL;HARD;3;CRITICAL - 4.5% (1475272 kB) free!
[1363695564] SERVICE ALERT: HHHHH;Current Load;CRITICAL;SOFT;2;CRITICAL - load average: 4.54, 19.91, 21.72
[1363695688] SERVICE ALERT: HHHHH;Current Load;WARNING;HARD;3;WARNING - load average: 2.68, 14.20, 19.33
[1363695806] SERVICE ALERT: HHHHH;Total Processes;OK;HARD;3;PROCS OK: 469 processes
[1363696158] SERVICE ALERT: HHHHH;Memory;WARNING;HARD;3;WARNING - 5.4% (1782328 kB) free!
[1363696287] SERVICE ALERT: HHHHH;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 601 processes
[1363696406] SERVICE ALERT: HHHHH;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 658 processes
[1363696524] SERVICE ALERT: HHHHH;Total Processes;OK;SOFT;2;PROCS OK: 422 processes
[1363696556] SERVICE ALERT: HHHHH;Disk usage;OK;HARD;3;DISK OK - free space: / 31976 MB (87% inode=99%): /home 3724 MB (39% inode=99%): /var 8916 MB (94% inode=99%): /usr 5810 MB (61% inode=93%): /u01 33786 MB (69% inode=97%): /archivelog 101039 MB (98% inode=99%): /data1 528852 MB (32% inode=99%): /redologs 49825 MB (97% inode=99%): /backup1 129122 MB (36% inode=99%): /index1 507012 MB (99% inode=99%): /data2 285003 MB (22% inode=99%):
[1363696668] SERVICE ALERT: HHHHH;Current Load;WARNING;SOFT;1;WARNING - load average: 5.47, 5.27, 5.07
[1363696759] SERVICE ALERT: HHHHH;Memory;CRITICAL;HARD;3;CRITICAL - 4.2% (1399672 kB) free!
[1363696784] SERVICE ALERT: HHHHH;Current Load;WARNING;SOFT;2;WARNING - load average: 5.56, 5.34, 5.12
[1363696887] SERVICE ALERT: HHHHH;Total Processes;WARNING;HARD;3;PROCS WARNING: 589 processes
[1363696887] SERVICE FLAPPING ALERT: HHHHH;Total Processes;STARTED; Service appears to have started flapping (23.2% change >= 20.0% threshold)
[1363696904] SERVICE ALERT: HHHHH;Current Load;WARNING;HARD;3;WARNING - load average: 5.19, 5.29, 5.13
[1363697164] HOST ALERT: HHHHH;DOWN;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363697234] HOST ALERT: HHHHH;UP;SOFT;2;NRPE v2.12
[1363697485] SERVICE ALERT: HHHHH;Current Load;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363697488] SERVICE ALERT: HHHHH;Current Load;OK;HARD;3;OK - load average: 0.89, 0.65, 3.06
[1363697504] SERVICE ALERT: HHHHH;Current Load;OK;HARD;3;OK - load average: 4.76, 4.81, 4.95
[1363697526] SERVICE ALERT: HHHHH;CPU Statistics;OK;HARD;3;CPU STATISTICS OK : user=69% system=31% iowait=0% idle=0% nice=0.00% steal=0.00%
[1363697565] SERVICE FLAPPING ALERT: HHHHH;Memory;STOPPED; Service appears to have stopped flapping (4.8% change < 5.0% threshold)
[1363697596] SERVICE ALERT: HHHHH;Current Load;OK;SOFT;2;OK - load average: 0.00, 0.00, 0.00
[1363697725] SERVICE ALERT: HHHHH;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 682 processes
[1363697844] SERVICE ALERT: HHHHH;Total Processes;OK;SOFT;2;PROCS OK: 378 processes
[1363698126] SERVICE ALERT: HHHHH;CPU Statistics;CRITICAL;SOFT;1;CPU STATISTICS CRITICAL : user=72% system=28% iowait=0% idle=0% nice=0.00% steal=0.00%
[1363698215] SERVICE ALERT: HHHHH;CPU Statistics;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363698245] SERVICE ALERT: HHHHH;CPU Statistics;CRITICAL;SOFT;2;CPU STATISTICS CRITICAL : user=72% system=28% iowait=0% idle=0% nice=0.00% steal=0.00%
[1363698325] SERVICE ALERT: HHHHH;CPU Statistics;OK;SOFT;2;CPU STATISTICS OK : user=44% system=1% iowait=0% idle=55%
[1363698364] SERVICE ALERT: HHHHH;CPU Statistics;CRITICAL;HARD;3;CPU STATISTICS CRITICAL : user=75% system=25% iowait=0% idle=0% nice=0.00% steal=0.00%
[1363698444] SERVICE ALERT: HHHHH;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 669 processes
[1363698558] SERVICE ALERT: HHHHH;Memory;WARNING;HARD;3;WARNING - 6.8% (2247752 kB) free!
[1363698564] SERVICE ALERT: HHHHH;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 667 processes
[1363698688] SERVICE ALERT: HHHHH;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 640 processes
[1363698688] SERVICE FLAPPING ALERT: HHHHH;Total Processes;STARTED; Service appears to have started flapping (22.5% change >= 20.0% threshold)
[1363699288] SERVICE ALERT: HHHHH;Total Processes;OK;HARD;3;PROCS OK: 382 processes
[1363699288] SERVICE ALERT: HHHHH;CPU Statistics;WARNING;SOFT;1;CPU STATISTICS WARNING : user=0% system=4% iowait=59% idle=37%
[1363699405] SERVICE ALERT: HHHHH;CPU Statistics;OK;SOFT;2;CPU STATISTICS OK : user=0% system=0% iowait=2% idle=98%
[1363699545] SERVICE ALERT: HHHHH;CPU Statistics;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363699554] HOST ALERT: HHHHH;DOWN;SOFT;1;Connection refused or timed out
[1363699604] SERVICE ALERT: HHHHH;CPU Statistics;CRITICAL;SOFT;2;Connection refused or timed out
[1363699625] HOST ALERT: HHHHH;DOWN;SOFT;2;Connection refused or timed out
[1363699625] SERVICE ALERT: HHHHH;PING;CRITICAL;HARD;1;Connection refused or timed out
[1363699665] SERVICE ALERT: HHHHH;CPU Statistics;CRITICAL;SOFT;3;Connection refused or timed out
[1363699685] SERVICE ALERT: HHHHH;Memory;CRITICAL;HARD;1;Connection refused or timed out
[1363699695] HOST ALERT: HHHHH;DOWN;SOFT;3;Connection refused or timed out
[1363699725] SERVICE ALERT: HHHHH;CPU Statistics;CRITICAL;HARD;4;Connection refused or timed out
[1363699727] SERVICE ALERT: HHHHH;Current Load;CRITICAL;HARD;1;Connection refused or timed out
[1363699727] SERVICE ALERT: HHHHH;Current Users;CRITICAL;HARD;1;Connection refused or timed out
[1363699764] HOST ALERT: HHHHH;DOWN;SOFT;4;Connection refused or timed out
[1363699786] SERVICE ALERT: HHHHH;Swap Usage;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363699834] HOST ALERT: HHHHH;UP;SOFT;5;NRPE v2.12
[1363699844] SERVICE ALERT: HHHHH;CPU Statistics;OK;HARD;4;CPU STATISTICS OK : user=0% system=0% iowait=0% idle=100%
[1363699888] SERVICE ALERT: HHHHH;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 682 processes
[1363699894] SERVICE ALERT: HHHHH;Swap Usage;OK;SOFT;2;SWAP OK - 100% free (1000 MB out of 1000 MB)
[1363700006] SERVICE ALERT: HHHHH;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 696 processes
[1363700124] SERVICE ALERT: HHHHH;Memory;WARNING;SOFT;1;WARNING - 8.2% (2688440 kB) free!
[1363700124] SERVICE ALERT: HHHHH;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 686 processes
[1363700204] SERVICE ALERT: HHHHH;Memory;CRITICAL;HARD;3;CRITICAL - 4.6% (745416 kB) free!
[1363700225] SERVICE ALERT: HHHHH;PING;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 0.13 ms
[1363700244] SERVICE ALERT: HHHHH;Memory;WARNING;SOFT;2;WARNING - 8.5% (2803768 kB) free!
[1363700284] SERVICE ALERT: HHHHH;Memory;OK;HARD;1;OK - 79.2% (52170160 kB) free.
[1363700325] SERVICE ALERT: HHHHH;Current Load;OK;HARD;1;OK - load average: 1.71, 1.28, 1.63
[1363700325] SERVICE ALERT: HHHHH;Current Users;OK;HARD;1;USERS OK - 2 users currently logged in
[1363700365] SERVICE ALERT: HHHHH;Memory;WARNING;HARD;3;WARNING - 8.8% (2902256 kB) free!
[1363700764] SERVICE ALERT: HHHHH;Total Processes;OK;HARD;3;PROCS OK: 486 processes
[1363700764] SERVICE FLAPPING ALERT: HHHHH;Total Processes;STARTED; Service appears to have started flapping (20.3% change >= 20.0% threshold)
[1363700805] SERVICE ALERT: HHHHH;Memory;WARNING;HARD;3;WARNING - 5.1% (834168 kB) free!
[1363701196] SERVICE ALERT: HHHHH;Current Load;OK;HARD;3;OK - load average: 4.56, 4.38, 4.91
[1363701224] SERVICE ALERT: HHHHH;Swap Usage;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363701224] SERVICE ALERT: HHHHH;Memory;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363701224] SERVICE ALERT: HHHHH;PING;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363701254] SERVICE ALERT: HHHHH;PING;WARNING;SOFT;1;PING WARNING - Packet loss = 0%, RTA = 104.45 ms
[1363701334] SERVICE ALERT: HHHHH;Swap Usage;OK;SOFT;2;SWAP OK - 100% free (16386 MB out of 16386 MB)
[1363701334] SERVICE ALERT: HHHHH;Memory;OK;SOFT;2;OK - 30.4% (4966144 kB) free.
[1363701335] SERVICE ALERT: HHHHH;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.27 ms
[1363701375] SERVICE ALERT: HHHHH;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 5.41 ms
[1363701545] SERVICE ALERT: HHHHH;Mysql Status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363701664] SERVICE ALERT: HHHHH;Mysql Status;OK;SOFT;2;check_mysql: OK (2 connections running)
[1363701915] SERVICE ALERT: HHHHH;User Processes Stats;WARNING;SOFT;1;CPU WARNING: 1 warn out of 168 processes
[1363701915] SERVICE ALERT: HHHHH;Total Processes;WARNING;SOFT;1;CPU WARNING: 1 warn out of 168 processes
[1363702034] SERVICE ALERT: HHHHH;User Processes Stats;WARNING;SOFT;2;CPU WARNING: 1 warn out of 164 processes
[1363702034] SERVICE ALERT: HHHHH;Total Processes;WARNING;SOFT;2;CPU WARNING: 1 warn out of 164 processes
[1363702157] SERVICE ALERT: HHHHH;User Processes Stats;WARNING;HARD;3;CPU WARNING: 1 warn out of 164 processes
[1363702159] SERVICE ALERT: HHHHH;Total Processes;WARNING;HARD;3;CPU WARNING: 1 warn out of 164 processes
[1363702164] SERVICE ALERT: HHHHH;Memory;OK;HARD;3;OK - 15.2% (4984536 kB) free.
[1363702454] SERVICE ALERT: HHHHH;Memory;WARNING;HARD;3;WARNING - 5.2% (3412776 kB) free!
[1363702454] SERVICE ALERT: HHHHH;Memory;WARNING;HARD;3;WARNING - 5.2% (3415808 kB) free!
[1363702455] SERVICE ALERT: HHHHH;Memory;WARNING;HARD;3;WARNING - 5.2% (3419248 kB) free!
[1363702524] SERVICE ALERT: HHHHH;Total Processes;WARNING;HARD;3;PROCS WARNING: 536 processes
[1363702554] SERVICE ALERT: HHHHH;Memory;WARNING;HARD;3;WARNING - 5.1% (1675288 kB) free!
[1363702757] SERVICE ALERT: HHHHH;User Processes Stats;OK;HARD;3;CPU OK: 163 processes
[1363702759] SERVICE ALERT: HHHHH;Total Processes;OK;HARD;3;CPU OK: 163 processes
[1363702894] SERVICE ALERT: HHHHH;Memory;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363703005] SERVICE ALERT: HHHHH;Memory;OK;SOFT;2;OK: Memory Usage - 30.03% RAM, 3.81% Swap
[1363703054] SERVICE ALERT: HHHHH;Memory;CRITICAL;HARD;3;CRITICAL - 4.8% (3131944 kB) free!
[1363703054] SERVICE ALERT: HHHHH;Memory;CRITICAL;HARD;3;CRITICAL - 4.8% (3134008 kB) free!
[1363703055] SERVICE ALERT: HHHHH;Memory;CRITICAL;HARD;3;CRITICAL - 4.8% (3133056 kB) free!
[1363703124] SERVICE ALERT: HHHHH;Total Processes;OK;HARD;3;PROCS OK: 480 processes
[1363703154] SERVICE ALERT: HHHHH;Memory;CRITICAL;HARD;3;CRITICAL - 4.6% (1501520 kB) free!
[1363703696] SERVICE ALERT: HHHHH;Memory;CRITICAL;HARD;3;CRITICAL - 4.3% (179064 kB) free!
[1363703804] SERVICE ALERT: HHHHH;Memory;CRITICAL;HARD;3;CRITICAL - 5.0% (815448 kB) free!
[1363704205] SERVICE ALERT: HHHHH;CPU Statistics;WARNING;SOFT;1;CPU STATISTICS WARNING : user=4% system=0% iowait=48% idle=48%
[1363704324] SERVICE ALERT: HHHHH;CPU Statistics;WARNING;SOFT;2;CPU STATISTICS WARNING : user=38% system=10% iowait=44% idle=8%
[1363704324] SERVICE ALERT: HHHHH;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 689 processes
[1363704444] SERVICE ALERT: HHHHH;CPU Statistics;WARNING;HARD;3;CPU STATISTICS WARNING : user=16% system=1% iowait=82% idle=0%
[1363704444] SERVICE ALERT: HHHHH;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 670 processes
[1363704559] SERVICE ALERT: HHHHH;Total Processes;OK;HARD;3;PROCS OK: 127 processes
[1363704564] SERVICE ALERT: HHHHH;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 677 processes
[1363704615] SERVICE ALERT: HHHHH;CPU Statistics;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363704734] SERVICE ALERT: HHHHH;CPU Statistics;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1363704854] SERVICE ALERT: HHHHH;CPU Statistics;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1363704994] SERVICE ALERT: HHHHH;Current Load;WARNING;SOFT;1;WARNING - load average: 16.54, 14.86, 9.31
[1363705114] SERVICE ALERT: HHHHH;Current Load;WARNING;SOFT;2;WARNING - load average: 13.31, 14.43, 9.84
[1363705164] SERVICE ALERT: HHHHH;Memory;WARNING;SOFT;1;WARNING - 10.0% (3289960 kB) free!
[1363705234] SERVICE ALERT: HHHHH;Current Load;WARNING;HARD;3;WARNING - load average: 5.49, 11.40, 9.31
[1363705290] SERVICE ALERT: HHHHH;Memory;WARNING;SOFT;2;WARNING - 10.0% (3285024 kB) free!
[1363705406] SERVICE ALERT: HHHHH;Memory;WARNING;HARD;3;WARNING - 10.0% (3270968 kB) free!
[1363705445] SERVICE ALERT: HHHHH;CPU Statistics;OK;HARD;3;OK: CPU=4.90
[1363705904] SERVICE ALERT: HHHHH;Swap Usage;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363706005] SERVICE ALERT: HHHHH;Memory;OK;HARD;3;OK - 10.3% (3364232 kB) free.
[1363706014] SERVICE ALERT: HHHHH;Swap Usage;OK;SOFT;2;SWAP OK - 100% free (1000 MB out of 1000 MB)
[1363706054] SERVICE ALERT: HHHHH;Memory;WARNING;HARD;3;WARNING - 5.1% (3353424 kB) free!
[1363706054] SERVICE ALERT: HHHHH;Memory;WARNING;HARD;3;WARNING - 5.1% (3355688 kB) free!
[1363706055] SERVICE ALERT: HHHHH;Memory;WARNING;HARD;3;WARNING - 5.1% (3358816 kB) free!
[1363706244] SERVICE FLAPPING ALERT: HHHHH;Memory;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1363706245] SERVICE ALERT: HHHHH;CPU Statistics;OK;HARD;3;CPU STATISTICS OK : user=41% system=3% iowait=8% idle=48%
[1363706434] SERVICE ALERT: HHHHH;Current Load;OK;HARD;3;OK - load average: 0.12, 1.38, 3.78
[1363706654] SERVICE ALERT: HHHHH;Memory;CRITICAL;HARD;3;CRITICAL - 4.8% (3127088 kB) free!
[1363706654] SERVICE FLAPPING ALERT: HHHHH;Memory;STARTED; Service appears to have started flapping (23.2% change >= 20.0% threshold)
[1363706654] SERVICE ALERT: HHHHH;Memory;CRITICAL;HARD;3;CRITICAL - 4.8% (3126000 kB) free!
[1363706654] SERVICE FLAPPING ALERT: HHHHH;Memory;STARTED; Service appears to have started flapping (23.2% change >= 20.0% threshold)
[1363706654] SERVICE ALERT: HHHHH;Memory;CRITICAL;HARD;3;CRITICAL - 4.8% (3129432 kB) free!
[1363706654] SERVICE FLAPPING ALERT: HHHHH;Memory;STARTED; Service appears to have started flapping (23.2% change >= 20.0% threshold)
[1363706875] SERVICE ALERT: HHHHH;Total Processes;OK;HARD;3;PROCS OK: 500 processes
[1363707125] SERVICE ALERT: HHHHH;Current Load;WARNING;SOFT;1;WARNING - load average: 27.02, 11.94, 6.15
[1363707245] SERVICE ALERT: HHHHH;Current Load;WARNING;SOFT;2;WARNING - load average: 24.33, 16.09, 8.46
[1363707364] SERVICE ALERT: HHHHH;Current Load;WARNING;HARD;3;WARNING - load average: 6.58, 12.21, 7.96
[1363707365] SERVICE FLAPPING ALERT: HHHHH;Total Processes;STOPPED; Service appears to have stopped flapping (4.8% change < 5.0% threshold)
[1363707404] SERVICE ALERT: HHHHH;Memory;WARNING;HARD;3;WARNING - 6.2% (1015568 kB) free!
[1363707404] SERVICE FLAPPING ALERT: HHHHH;Memory;STARTED; Service appears to have started flapping (21.2% change >= 20.0% threshold)
[1363707564] SERVICE ALERT: HHHHH;Total Processes;OK;HARD;3;PROCS OK: 428 processes
[1363707689] SERVICE FLAPPING ALERT: HHHHH;Current Load;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1363708164] SERVICE ALERT: HHHHH;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 673 processes
[1363708287] SERVICE FLAPPING ALERT: HHHHH;Total Processes;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1363708290] SERVICE ALERT: HHHHH;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 668 processes
[1363708396] SERVICE ALERT: HHHHH;Current Load;WARNING;SOFT;1;WARNING - load average: 7.01, 5.62, 5.08
[1363708405] SERVICE ALERT: HHHHH;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 698 processes
[1363708514] SERVICE ALERT: HHHHH;Current Load;WARNING;SOFT;2;WARNING - load average: 7.38, 6.30, 5.41
[1363708634] SERVICE ALERT: HHHHH;Current Load;WARNING;HARD;3;WARNING - load average: 6.01, 6.25, 5.50
[1363708665] SERVICE ALERT: HHHHH;Disk usage;WARNING;HARD;3;DISK WARNING - free space: / 53362 MB (73% inode=99%): /e2e 5113 MB (99% inode=99%): /obip 8070 MB (78% inode=99%): /oracle 1511 MB (29% inode=99%): /siebel1 14142 MB (15% inode=99%): /siebelfs 7693 MB (15% inode=98%):
[1363708758] SERVICE FLAPPING ALERT: HHHHH;Memory;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1363708764] SERVICE ALERT: HHHHH;Memory;OK;HARD;3;OK - 28.1% (19378600 kB) free.
[1363708794] SERVICE ALERT: HHHHH;Memory;OK;HARD;3;OK - 18.9% (13055288 kB) free.
[1363708797] SERVICE ALERT: HHHHH;Memory;OK;HARD;3;OK - 18.8% (12942016 kB) free.
[1363708904] SERVICE FLAPPING ALERT: HHHHH;Current Load;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1363709275] SERVICE ALERT: HHHHH;Total Processes;WARNING;SOFT;1;PROCS WARNING: 507 processes
[1363709276] SERVICE ALERT: HHHHH;CPU Statistics;WARNING;SOFT;1;CPU STATISTICS WARNING : user=1% system=1% iowait=51% idle=47%
[1363709364] SERVICE ALERT: HHHHH;Memory;WARNING;SOFT;1;WARNING - 7.4% (5079704 kB) free!
[1363709384] SERVICE ALERT: HHHHH;Current Load;WARNING;SOFT;1;WARNING - load average: 5.74, 5.77, 5.26
[1363709384] SERVICE ALERT: HHHHH;Memory;WARNING;SOFT;1;WARNING - 7.1% (4922376 kB) free!
[1363709384] SERVICE ALERT: HHHHH;Memory;WARNING;SOFT;1;WARNING - 7.1% (4925152 kB) free!
[1363709396] SERVICE ALERT: HHHHH;Total Processes;WARNING;SOFT;2;PROCS WARNING: 507 processes
[1363709400] SERVICE ALERT: HHHHH;CPU Statistics;OK;SOFT;2;CPU STATISTICS OK : user=1% system=6% iowait=20% idle=73%
[1363709489] SERVICE ALERT: HHHHH;Memory;WARNING;SOFT;2;WARNING - 6.4% (4404072 kB) free!
[1363709504] SERVICE ALERT: HHHHH;Current Load;WARNING;SOFT;2;WARNING - load average: 5.97, 5.78, 5.32
[1363709504] SERVICE ALERT: HHHHH;Memory;WARNING;SOFT;2;WARNING - 6.5% (4448024 kB) free!
[1363709504] SERVICE ALERT: HHHHH;Memory;WARNING;SOFT;2;WARNING - 6.5% (4450984 kB) free!
[1363709515] SERVICE ALERT: HHHHH;Total Processes;WARNING;HARD;3;PROCS WARNING: 508 processes
[1363709605] SERVICE ALERT: HHHHH;Memory;WARNING;HARD;3;WARNING - 5.8% (4003448 kB) free!
[1363709624] SERVICE ALERT: HHHHH;Current Load;WARNING;HARD;3;WARNING - load average: 6.05, 5.86, 5.40
[1363709625] SERVICE ALERT: HHHHH;Memory;WARNING;HARD;3;WARNING - 5.7% (3926200 kB) free!
[1363709626] SERVICE ALERT: HHHHH;Memory;WARNING;HARD;3;WARNING - 5.7% (3929144 kB) free!
[1363709844] SERVICE ALERT: HHHHH;Memory;WARNING;SOFT;1;WARNING - 9.8% (401312 kB) free!
[1363709964] SERVICE ALERT: HHHHH;Memory;WARNING;SOFT;2;WARNING - 9.8% (401296 kB) free!
[1363710086] SERVICE ALERT: HHHHH;Memory;WARNING;HARD;3;WARNING - 9.8% (401272 kB) free!
[1363710115] SERVICE ALERT: HHHHH;Total Processes;OK;HARD;3;PROCS OK: 500 processes
[1363710205] SERVICE ALERT: HHHHH;Total Processes;OK;HARD;3;PROCS OK: 424 processes
[1363710205] SERVICE ALERT: HHHHH;Memory;CRITICAL;HARD;3;CRITICAL - 3.4% (2344896 kB) free!
[1363710224] SERVICE ALERT: HHHHH;Memory;CRITICAL;HARD;3;CRITICAL - 3.5% (2420000 kB) free!
[1363710225] SERVICE ALERT: HHHHH;Memory;CRITICAL;HARD;3;CRITICAL - 3.5% (2422984 kB) free!
[1363710404] SERVICE ALERT: HHHHH;Memory;CRITICAL;HARD;3;CRITICAL - 4.9% (801000 kB) free!
[1363710444] SERVICE ALERT: HHHHH;Memory;OK;HARD;3;OK - 66.3% (5410040 kB) free.
[1363710465] SERVICE ALERT: HHHHH;Disk usage;CRITICAL;HARD;3;DISK CRITICAL - free space: / 53355 MB (73% inode=99%): /e2e 5113 MB (99% inode=99%): /obip 8070 MB (78% inode=99%): /oracle 1511 MB (29% inode=99%): /siebel1 10041 MB (10% inode=98%): /siebelfs 7693 MB (15% inode=98%):
[1363710714] SERVICE ALERT: HHHHH;Total Processes;WARNING;SOFT;1;PROCS WARNING: 517 processes
[1363710806] SERVICE ALERT: HHHHH;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 608 processes
[1363710835] SERVICE ALERT: HHHHH;Total Processes;WARNING;SOFT;2;PROCS WARNING: 514 processes
[1363710924] SERVICE ALERT: HHHHH;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 653 processes
[1363710955] SERVICE ALERT: HHHHH;Total Processes;WARNING;HARD;3;PROCS WARNING: 515 processes
[1363710955] SERVICE FLAPPING ALERT: HHHHH;Total Processes;STARTED; Service appears to have started flapping (23.8% change >= 20.0% threshold)
[1363711024] HOST ALERT: HHHHH;DOWN;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363711044] SERVICE ALERT: HHHHH;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 665 processes
[1363711055] SERVICE ALERT: HHHHH;Memory;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363711104] HOST ALERT: HHHHH;DOWN;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1363711124] HOST ALERT: HHHHH;UP;SOFT;3;NRPE v2.10
[1363711199] SERVICE ALERT: HHHHH;CPU Statistics;WARNING;SOFT;1;CPU STATISTICS WARNING : user=0% system=0% iowait=51% idle=48%
[1363711289] SERVICE ALERT: HHHHH;Current Load;CRITICAL;SOFT;1;CRITICAL - load average: 35.13, 33.60, 21.76
[1363711315] SERVICE ALERT: HHHHH;CPU Statistics;WARNING;SOFT;2;CPU STATISTICS WARNING : user=1% system=1% iowait=36% idle=63%
[1363711405] SERVICE ALERT: HHHHH;Current Load;CRITICAL;SOFT;2;CRITICAL - load average: 28.40, 31.63, 22.50
[1363711435] SERVICE ALERT: HHHHH;CPU Statistics;OK;SOFT;3;CPU STATISTICS OK : user=1% system=10% iowait=16% idle=72%
[1363711524] SERVICE ALERT: HHHHH;Current Load;CRITICAL;HARD;3;CRITICAL - load average: 28.38, 30.36, 23.15
[1363711644] SERVICE ALERT: HHHHH;CPU Statistics;WARNING;SOFT;1;CPU STATISTICS WARNING : user=2% system=2% iowait=72% idle=24%
[1363711644] SERVICE ALERT: HHHHH;Memory;CRITICAL;SOFT;1;CRITICAL - 1.5% (124096 kB) free!
[1363711764] SERVICE ALERT: HHHHH;Memory;CRITICAL;SOFT;2;CRITICAL - 1.5% (123552 kB) free!
[1363711764] SERVICE ALERT: HHHHH;CPU Statistics;OK;SOFT;2;CPU STATISTICS OK : user=0% system=2% iowait=21% idle=77%
[1363711888] SERVICE ALERT: HHHHH;Memory;CRITICAL;HARD;3;CRITICAL - 2.0% (166808 kB) free!
[1363712125] SERVICE ALERT: HHHHH;Current Load;WARNING;HARD;3;WARNING - load average: 0.08, 6.08, 13.91
[1363712204] SERVICE ALERT: HHHHH;Memory;WARNING;HARD;3;WARNING - 5.2% (857808 kB) free!
[1363712804] SERVICE ALERT: HHHHH;Memory;CRITICAL;HARD;3;CRITICAL - 3.2% (522280 kB) free!
[1363713324] SERVICE ALERT: HHHHH;Current Load;OK;HARD;3;OK - load average: 0.02, 0.19, 3.73
[1363713536] SERVICE ALERT: HHHHH;Memory;CRITICAL;HARD;3;CRITICAL - 2.0% (160192 kB) free!
[1363713835] SERVICE ALERT: HHHHH;CPU Statistics;WARNING;SOFT;1;CPU STATISTICS WARNING : user=5% system=17% iowait=30% idle=48%
[1363713955] SERVICE ALERT: HHHHH;CPU Statistics;OK;SOFT;2;CPU STATISTICS OK : user=6% system=20% iowait=27% idle=46%
[1363714034] SERVICE ALERT: HHHHH;Current Load;OK;HARD;3;OK - load average: 3.45, 3.55, 4.41
[1363714136] SERVICE ALERT: HHHHH;Memory;WARNING;HARD;3;WARNING - 2.1% (172896 kB) free!
[1363714524] SERVICE ALERT: HHHHH;CPU Statistics;WARNING;SOFT;1;CPU STATISTICS WARNING : user=23% system=24% iowait=53% idle=0%
[1363714644] SERVICE ALERT: HHHHH;CPU Statistics;OK;SOFT;2;CPU STATISTICS OK : user=45% system=29% iowait=26% idle=0%
[1363715085] SERVICE ALERT: HHHHH;Swap Usage;UNKNOWN;SOFT;1;CHECK_NRPE: Received 0 bytes from daemon.  Check the remote server logs for error messages.
[1363715164] SERVICE ALERT: HHHHH;CPU Statistics;OK;HARD;3;CPU STATISTICS OK : user=69% system=31% iowait=0% idle=0% nice=0.00% steal=0.00%
[1363715204] SERVICE ALERT: HHHHH;Swap Usage;OK;SOFT;2;SWAP OK - 100% free (8177 MB out of 8191 MB)
[1363715244] SERVICE ALERT: HHHHH;CPU Statistics;WARNING;SOFT;1;CPU STATISTICS WARNING : user=36% system=23% iowait=37% idle=3%
[1363715364] SERVICE ALERT: HHHHH;CPU Statistics;WARNING;SOFT;2;CPU STATISTICS WARNING : user=14% system=24% iowait=62% idle=0%
[1363715490] SERVICE ALERT: HHHHH;CPU Statistics;WARNING;HARD;3;CPU STATISTICS WARNING : user=30% system=24% iowait=34% idle=12%
[1363715764] SERVICE ALERT: HHHHH;CPU Statistics;CRITICAL;SOFT;1;CPU STATISTICS CRITICAL : user=82% system=18% iowait=0% idle=0% nice=0.00% steal=0.00%
[1363715804] SERVICE ALERT: HHHHH;Memory;WARNING;HARD;3;WARNING - 5.5% (895448 kB) free!
[1363715884] SERVICE ALERT: HHHHH;CPU Statistics;CRITICAL;SOFT;2;CPU STATISTICS CRITICAL : user=83% system=17% iowait=0% idle=0% nice=0.00% steal=0.00%
[1363716007] SERVICE ALERT: HHHHH;CPU Statistics;CRITICAL;HARD;3;CPU STATISTICS CRITICAL : user=81% system=19% iowait=0% idle=0% nice=0.00% steal=0.00%
[1363716555] SERVICE ALERT: HHHHH;Disk usage;CRITICAL;HARD;3;DISK CRITICAL - free space: / 2814 MB (70% inode=93%): /data1 3840 MB (38% inode=99%): /u01 2714 MB (27% inode=90%): /u05 6523 MB (19% inode=99%): /u03 1667 MB (16% inode=98%): /u06 3399 MB (10% inode=99%): /u02 5248 MB (52% inode=99%): /data2 5386 MB (53% inode=99%): /u04 3829 MB (38% inode=99%): /export/home 2716 MB (91% inode=99%):
[1363716577] SERVICE ALERT: HHHHH;Memory;WARNING;HARD;3;WARNING - 5.4% (1761344 kB) free!
[1363716691] SERVICE ALERT: HHHHH;CPU Statistics;OK;HARD;3;CPU STATISTICS OK : user=46% system=26% iowait=28% idle=0%
[1363716825] SERVICE ALERT: HHHHH;Current Load;OK;HARD;3;OK - load average: 3.18, 3.46, 4.48
[1363717044] SERVICE ALERT: HHHHH;Total Processes;OK;HARD;3;PROCS OK: 472 processes
[1363717164] SERVICE ALERT: HHHHH;CPU Statistics;WARNING;SOFT;1;CPU STATISTICS WARNING : user=48% system=6% iowait=46% idle=0%
[1363717184] SERVICE ALERT: HHHHH;Disk usage;CRITICAL;SOFT;1;DISK CRITICAL - free space: / 25349 MB (74% inode=98%): /u01 22442 MB (46% inode=97%): /data1 72892 MB (30% inode=99%): /data2 107214 MB (44% inode=99%): /archivelog 38093 MB (99% inode=99%): /redolog 26188 MB (91% inode=99%): /backup1 5673 MB (5% inode=99%): /app01 33309 MB (99% inode=99%):
[1363717288] SERVICE ALERT: HHHHH;CPU Statistics;WARNING;SOFT;1;CPU STATISTICS WARNING : user=20% system=24% iowait=56% idle=0%
[1363717288] SERVICE ALERT: HHHHH;CPU Statistics;OK;SOFT;2;CPU STATISTICS OK : user=60% system=40% iowait=0% idle=0%
[1363717304] SERVICE ALERT: HHHHH;Disk usage;CRITICAL;SOFT;2;DISK CRITICAL - free space: / 25349 MB (74% inode=98%): /u01 22442 MB (46% inode=97%): /data1 72892 MB (30% inode=99%): /data2 107214 MB (44% inode=99%): /archivelog 38093 MB (99% inode=99%): /redolog 26188 MB (91% inode=99%): /backup1 4663 MB (4% inode=99%): /app01 33309 MB (99% inode=99%):
[1363717405] SERVICE ALERT: HHHHH;CPU Statistics;WARNING;SOFT;2;CPU STATISTICS WARNING : user=24% system=31% iowait=45% idle=0%
[1363717406] SERVICE FLAPPING ALERT: HHHHH;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1363717424] SERVICE ALERT: HHHHH;Disk usage;OK;SOFT;3;DISK OK - free space: / 25349 MB (74% inode=98%): /u01 22442 MB (46% inode=97%): /data1 72892 MB (30% inode=99%): /data2 107214 MB (44% inode=99%): /archivelog 38093 MB (99% inode=99%): /redolog 26188 MB (91% inode=99%): /backup1 13108 MB (13% inode=99%): /app01 33309 MB (99% inode=99%):
[1363717524] SERVICE ALERT: HHHHH;CPU Statistics;OK;SOFT;3;CPU STATISTICS OK : user=47% system=26% iowait=27% idle=0%
[1363717644] SERVICE ALERT: HHHHH;Total Processes;WARNING;SOFT;1;PROCS WARNING: 526 processes
[1363717685] SERVICE ALERT: HHHHH;Current Load;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363717704] HOST ALERT: HHHHH;DOWN;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363717725] SERVICE ALERT: HHHHH;SMTP;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1363717764] SERVICE ALERT: HHHHH;Total Processes;OK;SOFT;2;PROCS OK: 458 processes
[1363717774] HOST ALERT: HHHHH;UP;SOFT;2;NRPE v2.12
[1363717777] SERVICE ALERT: HHHHH;Memory;CRITICAL;HARD;3;CRITICAL - 5.0% (1650592 kB) free!
[1363717798] SERVICE ALERT: HHHHH;Current Load;OK;SOFT;2;OK - load average: 4.54, 3.87, 2.29
[1363717800] HOST DOWNTIME ALERT: HHHHH;STARTED; Host has entered a period of scheduled downtime
[1363718100] HOST DOWNTIME ALERT: HHHHH;STOPPED; Host has exited from a period of scheduled downtime
[1363718200] SERVICE DOWNTIME ALERT: HHHHH;HTTP;STARTED; Service has entered a period of scheduled downtime
[1363718500] SERVICE DOWNTIME ALERT: HHHHH;HTTP;STOPPED; Service has exited from a period of scheduled downtime
[1363717800] HOST DOWNTIME ALERT: HHHHH;STARTED; Host has entered a period of scheduled downtime
[1363718100] HOST DOWNTIME ALERT: HHHHH;STOPPED; Host has exited from a period of scheduled downtime
[1363718200] SERVICE DOWNTIME ALERT: HHHHH;HTTP;STARTED; Service has entered a period of scheduled downtime
[1363718500] SERVICE DOWNTIME ALERT: HHHHH;HTTP;STOPPED; Service has exited from a period of scheduled downtime
[1363717800] HOST DOWNTIME ALERT: HHHHH;STARTED; Host has entered a period of scheduled downtime
[1363718100] HOST DOWNTIME ALERT: HHHHH;STOPPED; Host has exited from a period of scheduled downtime
[1363718200] SERVICE DOWNTIME ALERT: HHHHH;HTTP;STARTED; Service has entered a period of scheduled downtime
[1363718500] SERVICE DOWNTIME ALERT: HHHHH;HTTP;STOPPED; Service has exited from a period of scheduled downtime
[1363717800] HOST DOWNTIME ALERT: HHHHH;STARTED; Host has entered a period of scheduled downtime
[1363718100] HOST DOWNTIME ALERT: HHHHH;STOPPED; Host has exited from a period of scheduled downtime
[1363718200] SERVICE DOWNTIME ALERT: HHHHH;HTTP;STARTED; Service has entered a period of scheduled downtime
[1363718500] SERVICE DOWNTIME ALERT: HHHHH;HTTP;STOPPED; Service has exited from a period of scheduled downtime